```python
# Initialize Otter
import otter
grader = otter.Notebook("lab03.ipynb")
```

# E7: Lab Assignment 03

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to familiarize you with basic Python data structures.

## Autograder
You will be provided with some test cases (i.e., sample input data) and associated answers (i.e., expected outputs) that you can use to help check your code. The provided test cases are not exhaustive, and it is your responsibility to ensure that your code works in general, not just for a few supplied test cases. We may use additional hidden test cases in grading your lab assignments.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Your answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
Question  | Points
:---      | --:
1.0 – 1.3 | 4.0
2.0 – 2.1 | 4.0
3.0 – 3.8 | 9.0
4         | 4.0
5.0 | 3.0
5.1 | 1.0
Total | 25.0

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import numpy as np
import math
import matplotlib.pyplot as plt
import re
np.seterr(all='ignore');
import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Introduction
For this lab assignment, you will analyze water storage of three reservoirs in California from Jan-Dec 2022. First, you will clean and manipulate the data and then, you will calculate basic statistics for each reservoir.

### Background
The three reservoirs you will analyze are Lake Oroville, Shasta Lake, and Hetch Hetchy Reservoir. Shasta Lake and Lake Oroville are the largest and second-largest reservoirs in California, providing flood-control, hydropower, and water storage for the state. Shasta Lake is a key reservoir within the federal government's Central Valley Project. Similarly, Lake Oroville is a major reservoir in the state government's California State Water Project. Hetch Hetchy reservoir is within Yosemite National Park and provides water to the City of San Francisco and other counties in the San Francisco Bay Area.

### Data
We will read a text file that includes the monthly storage volumes for each reservoir, taken from January through December 2022, in units of acre-feet. We will provide you with the command to read the file, but in the future, you will learn how to read text files and other file types. In the code cell below, we are reading a text file and saving its values as `reservoir_data`. Run the cell below and do not modify its contents.

```python
# Please run this cell, and do not modify the contents
reservoir_data = np.loadtxt('resources/res_vals_int.csv',delimiter='\t')
print(reservoir_data)
```

The 12 rows correspond to the monthly storage volumes and the three columns correspond to the three reservoirs in the following order:
1. Lake Oroville
2. Shasta Lake
3. Hetch Hetchy reservoir

## Question 1.0

Create a list `reservoir_names` that includs the following two elements: `'Lake Oroville'` and `'Shasta'` (in the same order as listed here).

```python
# ANSWER CELL
...
```

```python
# Do not edit this cell. This exists for grading purposes.
q1_0 = reservoir_names.copy()
```

```python
grader.check("q1.0")
```

## Question 1.1

In the list `reservoir_names`, Shasta Lake is listed as just `'Shasta'`. Change the value from `'Shasta'` to `'Shasta Lake'`. Do not create a new variable `reservoir_names` – use list indexing.

```python
# ANSWER CELL
...
```

```python
# Do not edit this cell. This exists for grading purposes.
q1_1 = reservoir_names.copy()
```

```python
grader.check("q1.1")
```

## Question 1.2
Our list `reservoir_names` only has two entries! We are missing the third entry `'Hetch Hetchy Reservoir'`. Add `'Hetch Hetchy Reservoir'` to the end of the `reservoir_names` list. Do not redefine the list. Use appropriate list methods that allow you to add an entry to an existing list.

```python
# ANSWER CELL
...
```

```python
# Do not edit this cell. This exists for grading purposes.
q1_2 = reservoir_names.copy()
```

```python
grader.check("q1.2")
```

## Question 1.3
Now that you have the correct values in `reservoir_names`, check to see if certain reservoirs are in your list. Check if `"Pardee Reservoir"` is in the list of reservoirs using list methods. Write your boolean statement and assign it to `q1_3`.

*Note:* Pardee Reservoir is a major reservoir for East Bay Municipal Utility District.

```python
# ANSWER CELL
q1_3 = ...
```

```python
grader.check("q1.3")
```

## Question 2
In this question, you will create a range of numbers 1-12 to represent the rows within the `reservoir_data` array (i.e., the months Jan - Dec).

## Question 2.0
Create range of integers from 1 to 12, inclusive, using the `range` function. Store the range of numbers as `months_range`.

```python
# ANSWER CELL
...
```

```python
grader.check("q2.0")
```

## Question 2.1
Convert `months_range` to a numpy array and to a list. Store the values as `months_nparray` and `months_list` respectively. The variables `months_nparray` and `months_list` should contain the integers 1 through 12, just like `months_range`.

```python
# ANSWER CELL
months_nparray = ...
months_list = ...
print(months_nparray , months_list)
```

```python
grader.check("q2.1")
```

## Question 3
Now you will start to analyze the reservoir data. For this question, do not edit `reservoir_data`. Refer to the instructions at the beginning of the assignment on which column of the data corresponds to which reservoir.

## Question 3.0
First, slice the array to select the column of data corresponding to Lake Oroville.

Store the sliced data as `OrovilleData`.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.0")
```

## Question 3.1
Slice the array to select the winter storage volumes for Hetch Hetchy Reservoir. For this question, use January, February, and March as the winter months.

Store the sliced data as `WinterHetchHetchy`.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.1")
```

## Question 3.2
Slice the array to select the storage volumes for the evens months (February, April, ..., October, and December) for Lake Oroville and Hetch Hetchy Reservoir only.

Store the sliced data as `EvenMonths`.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.2")
```

## Question 3.3
For each month, how much more water was stored in Lake Oroville than in Hetch Hetchy? Report your answer in thousand acre-feet (TAF).

Store the differences in thousand acre-feet as `Diff`.

*Hint*: Your answer should be an array which includes the storage difference for each month. This can be achieved using arithmetic operations on slices of the array. Divide by 1000 to convert from acre-feet to TAF.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.3")
```

## Question 3.4
Find the mean storage volume for Summer (July through September) for Shasta Lake. California receives very little rainfall in the summer months (July through September), so reservoir storage during these months is of particular interest.

Store the mean value as `SummerMeanStorage`.

*Hint*: Slice the data to the months and reservoir of interest, then calculate the mean of the sliced array.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.4")
```

## Question 3.5
Check if the the mean value you computed above, `SummerMeanStorage`, exists in `reservoir_data`. Write your boolean statement and assign it to `q3_5`.

```python
# ANSWER CELL
q3_5 = ...
```

```python
grader.check("q3.5")
```

## Question 3.6
For each month, did the storage volume in **Lake Oroville and Shasta Lake** exceed `SummerMeanStorage`?

You should answer this by creating a logical expression using the array `reservoir_data`. The output should be another array that includes `True` or `False` for each month, depending on whether each month had storages higher than `SummerMeanStorage` for both Lake Oroville and Shasta Lake. This should be stored in `q3_6`.

```python
# ANSWER CELL
q3_6 = ...
```

```python
grader.check("q3.6")
```

## Question 3.7
For how many months did the storage volume in **Lake Oroville and Shasta Lake** exceed `SummerMeanStorage`?

Store the value as `q3_7`.

*Hint:* True is equivalent to 1 and False is equivalent to 0. So, you can take the sum of the logical array from the previous question to find the total number of months for which the logical expression is True.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.7")
```

## Question 3.8
We would like to convert the data in `reservoir_data` from acre-feet to $m^3$. 1 acre-feet is equal to 1233.48 $m^3$.

Store the value as `reservoir_data_m3`.

```python
# ANSWER CELL
reservoir_data_m3 = ...
```

```python
grader.check("q3.8")
```

## Question 4
Write a function called `ReservoirInfo` which takes information about reservoir storage as the inputs and outputs a statement: ‘In [Month Name], [Reservoir Name] held [Volume] acre-feet of water.’ where the words in brackets are the name of the input month, reservoir name, and storage volume, respectively. The function arguments should be:

* `name`: reservoir's name as a string, exactly as it appears in `reservoir_names`
* `month`: integer representing the month number, where 1 corresponds to January, 2 corresponds to February, etc.

Be careful with your string formatting. Any small modification, such as extra space or character, will result in a failed test.

You should name the function arguments `name` and `month`. An example function call and its output are:

```
>>>ReservoirInfo(name = 'Lake Oroville', month = 1)
'In January, Lake Oroville held 1642000 acre-feet of water.'
```

Important notes:
* Your function should **return** the statement
* [Month Name] should be the full month name of the input number. You may find the [`calendar.month_name[number]`](https://docs.python.org/3/library/calendar.html#calendar.month_name) attribute helpful. To use it, you have to import the module `calendar` first.
* [Volume] should be obtained from `reservoir_data` based on the corresponding input month and reservoir name. You may find the list `index()` method helpful (returns the index of the specified element in the list). Note that [Volume] should be an **integer** and not float.

```python
# ANSWER CELL
...
```

```python
grader.check("q4")
```

## Question 5.0
Finally for this assignment, you will define the input data and title string for a pre-made plotting function called `plot_storage`. The `plot_storage` function has two arguments:
* `name`: reservoir's name as a string, exactly as it appears in `reservoir_names`
* `data` is an array containing the storage data of the reservoir `name` (for only one reservoir)

Inside the function `plot_storage`, define `title_string` as a format string to have the reservoir name input, `name`, followed by "storage in 2022". For example, a graph of Lake Oroville should have the title "Lake Oroville storage in 2022". The variable `title_string` should automatically update based on the input `name`. Be careful with your string formatting. Any small modification, such as extra space or character, will result in a failed test.

Additionally, inside the function `plot_storage`, define `x_label` and `y_label` as format strings for the axes labels of your plot. Your labels should include the correct units "month" and "acre-feet". Be careful with your string formatting. Any misspelling will result in a failed test.

Once you have updated the missing statements in `plot_storage`, you will plot Hetch Hetchy data using the `plot_storage` function.
Define `name_input` as the value from the `reservoir_names` list corresponding to Hetch Hetchy. 
Define `data_input` as an array slice from `reservoir_data` corresponding to Hetch Hetchy storage data.

Finally, call the function `plot_storage` and pass the parameters `name_input` and `data_input`. Assign the function call to `q5_0`.

*Note:* This question relies on the completion of Q1.2 and Q2.0.

```python
# ANSWER CELL
def plot_storage(name, data):
    title_string = ...
    x_label = ...
    y_label = ...
    
    # Do not edit the function code below this line.
    plt.plot(months_range, data)
    plt.xlabel(x_label)
    plt.ylabel(y_label)
    plt.title(title_string)
    plt.show()
    return title_string , x_label , y_label

name_input = ...
data_input = ...

# # Assign the function call to q5_0
q5_0 = ...
```

```python
grader.check("q5.0")
```

## Question 5.1
In Q5.0, all three outputs from the `plot_storage` function were saved as a single variable called `q5_0`. 

Evaluate the type of `q5_0` and save the type as `q5_1a`.\
Evaluate the length of `q5_0` and save the length as `q5_1b`.

```python
# ANSWER CELL
q5_1a = ...
q5_1b = ...
print(f'Type: {q5_1a} \nLength: {q5_1b}')
```

```python
grader.check("q5.1")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 3!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 1)
```
>>> assert get_hash(len(q1_0)) == 'c81e728d9d4c2f636f067f89cc14862c' , 'You should only have two values (for now).'
>>> assert get_hash(q1_0[0]) == 'b947db2fb4b5855984587a0d4fabf896' , 'Check value of Lake Oroville entry'
>>> assert get_hash(q1_0[1]) == '36d4e7e1bfd840adcaede22706f227f9' , 'Check value of Shasta Lake entry'
>>> assert get_hash(type(q1_0)) == '425ed5b515f0a67c6316f0ec499f86bf', 'Check data type of reservoir_names'
```


Tests for q1.1 (points: 1)
```
>>> assert get_hash(len(q1_1)) == 'c81e728d9d4c2f636f067f89cc14862c' , 'You should still only have two values.'
>>> assert get_hash(q1_1[0]) == 'b947db2fb4b5855984587a0d4fabf896' , 'Do not edit the Lake Oroville entry'
>>> assert get_hash(q1_1[1]) =='dd7d39986c5d807a0248838b12acc7f3' , 'Check value of Shasta Lake entry'
```


Tests for q1.2 (points: 1)
```
>>> assert get_hash(len(q1_2)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of list'
>>> assert get_hash(q1_2[-1].upper()) == 'e5bb3ed8fc109fa9f05b0d538a30edf8' , 'Check value Hetch Hetchy entry'
>>> assert get_hash(q1_2[-1]) == 'f87798ca0a0f836f16995cc5c7955c2a' , 'Check value Hetch Hetchy entry'
```


Tests for q1.3 (points: 1)
```
>>> assert get_hash(q1_3) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q2.0 (points: 2)
```
>>> assert get_hash(months_range) == '5319bc85f3a29d6cd7b6b65df958ae78'
>>> assert get_hash(months_range[0]) == 'c4ca4238a0b923820dcc509a6f75849b' , 'Check start value'
>>> assert get_hash(months_range[-1]) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check end value'
>>> assert get_hash(sum(months_range)) == '35f4a8d465e6e1edc05f3d8ab658c551' , 'Check step size'
```


Tests for q2.1 (points: 2)
```
>>> assert get_hash(type(months_nparray)) == '30b39aca73d0815f5ca1a06af1412947'
>>> assert get_hash(len(months_nparray)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_nparray)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```
```
>>> assert get_hash(type(months_list)) == '425ed5b515f0a67c6316f0ec499f86bf'
>>> assert get_hash(len(months_list)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(sum(months_list)) == '35f4a8d465e6e1edc05f3d8ab658c551'
```


Tests for q3.0 (points: 1)
```
>>> assert get_hash(len(OrovilleData)) == 'c20ad4d76fe97759aa27a0c99bff6710' , 'Check length of selection'
>>> assert get_hash(int(sum(OrovilleData))) == '97c2a51b4ddc77196fa563cd72957c1e' , 'Check which column corresponds to Lake Oroville'
>>> assert get_hash(OrovilleData.astype(int)) == 'd57f3682934eebd41c73b68817c6592d'
```


Tests for q3.1 (points: 1)
```
>>> assert get_hash(len(WinterHetchHetchy)) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3' , 'Check length of selection'
>>> assert get_hash(int(sum(WinterHetchHetchy))) == '50037bd6578a3e4db86ac430e190f7cb' , 'Check which column corresponds to Lake Oroville'
>>> assert get_hash(WinterHetchHetchy.astype(int)) == '7b67ebf4e3d0d34b2c62dba86b620e1b'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(len(EvenMonths)) == '1679091c5a880faf6fb5e6087eb1b2dc' , 'Check length of selection'
>>> assert get_hash(EvenMonths.shape) == 'e186b2fc415e1b6085ca473fcc974c66' , 'Check slice selection'
>>> assert get_hash(EvenMonths.astype(int)) == 'b2f43e102ce4fe6ac1033f199ea9cd4b'
```


Tests for q3.3 (points: 1)
```
>>> assert get_hash(len(Diff)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(Diff.shape) == '3f30626b8bdbdc8366257c053839456f'
>>> assert get_hash(sum(Diff)) == '4a3dd518aa065cd57d5d5acfeb1ff6da'
>>> assert get_hash(Diff) == 'bde6ec575d6f776bfcb4df36f0e64c28'
```


Tests for q3.4 (points: 1)
```
>>> assert get_hash(round(SummerMeanStorage)) == '118df32ac7db6de47c75d4ac84e6af2e' , 'Check mean calculation'
```


Tests for q3.5 (points: 1)
```
>>> assert get_hash(q3_5) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.6 (points: 1)
```
>>> assert get_hash(len(q3_6)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(q3_6.shape) == '3f30626b8bdbdc8366257c053839456f'
>>> assert get_hash(q3_6) == '453afeb97740796e13cd9c391a6c9c86'
```


Tests for q3.7 (points: 1)
```
>>> assert get_hash(q3_7) == '1679091c5a880faf6fb5e6087eb1b2dc'
```


Tests for q3.8 (points: 1)
```
>>> assert get_hash(len(reservoir_data_m3)) == 'c20ad4d76fe97759aa27a0c99bff6710'
>>> assert get_hash(reservoir_data_m3.shape) == '52658b9c1506282ca4f20ade14e21d99'
>>> assert get_hash(reservoir_data_m3.astype(int)) == 'dc0b8fe95b9cb8a9b53140593a64d969'
```


Tests for q4 (points: 4)
```
>>> # Check data structure
>>> import calendar
>>> assert get_hash(len(ReservoirInfo(name=reservoir_names[2],month=2))) == '735b90b4568125ed6c3f678819b6e058' , 'Check spacing and data type'
```
```
>>> # Check the return type
>>> assert get_hash(type(ReservoirInfo(name=reservoir_names[1],month=5).upper())) == '9a86641cdf2fdb47f786dc955264738d', 'Check the return type'
```
```
>>> # Check wording of string output
>>> assert get_hash("HELD" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the wording'
>>> assert get_hash("WATER" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the wording'
```
```
>>> # Check reservoir name
>>> assert get_hash("SHASTA" in ReservoirInfo(name=reservoir_names[1],month=8).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the lake name'
>>> assert get_hash("OROVILLE" in ReservoirInfo(name=reservoir_names[0],month=8).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the lake name'
```
```
>>> # Check the name of the month
>>> assert get_hash("MAY" in ReservoirInfo(name=reservoir_names[1],month=5).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the month name'
>>> assert get_hash("JULY" in ReservoirInfo(name=reservoir_names[1],month=7).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the month name'
>>> assert get_hash("FEBRUARY" in ReservoirInfo(name=reservoir_names[1],month=2).upper()) == 'f827cf462f62848df37c5e1e94a4da74', 'Check the month name'
```
```
>>> # Check value
>>> assert get_hash("1621000" in ReservoirInfo(name=reservoir_names[1],month=1)) == 'f827cf462f62848df37c5e1e94a4da74', 'Check value'
>>> assert get_hash("1690000" in ReservoirInfo(name=reservoir_names[1],month=2)) == 'f827cf462f62848df37c5e1e94a4da74', 'Check value'
>>> assert get_hash("305100" in ReservoirInfo(name=reservoir_names[2],month=8)) == 'f827cf462f62848df37c5e1e94a4da74', 'Check value'
```


Tests for q5.0 (points: 3)
```
>>> assert get_hash(name_input) == 'f87798ca0a0f836f16995cc5c7955c2a'
>>> assert get_hash(q5_0[0]) == '75cbfe9a1d825c40fd979cc5614764aa' , 'Check title string'
>>> assert get_hash("MONTH" in q5_0[1].upper()) == 'f827cf462f62848df37c5e1e94a4da74' , 'Check x-label'
>>> # Uses regex search to look for acre feet. Any character allowed between acre and feet. Case insensitive.
>>> assert re.search(r'ACRE.FEET',q5_0[2].upper()) , 'Check y-label'
>>> assert get_hash(data_input.astype(int)) == 'ac3c5f2b1ab399ae627808b2459456ca'
```


Tests for q5.1 (points: 1)
```
>>> assert get_hash(q5_1a) == '22b3afce62075c7012f8e5041adfee16'
>>> assert get_hash(q5_1b) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
```
