```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# E7: Lab Assignment 02

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to familiarize you with functions.

## Autograder
For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee that it will work for all possible test cases relevant to the question. It is your responsibility to test your function thoroughly, to ensure that it will also work in situations not covered by the test cases provided.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. `...` is just a placeholder.

You answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

If you want to read documentation on a Python function, you can type `help()` or `help(function)`.


## Score Breakdown
Question| Points
--- | --
1 | 4
2.0 | 3
2.1 | 3
2.2 | 3
2.3 | 3
3.0 | 3
3.1 | 3
3.2 | 1
3.3 | 2
Total | 25

```python
# Please run this cell, and do not modify the contets
import math
import numpy as np
import matplotlib.pyplot as plt
#plt.style.use('fivethirtyeight')
#plt.rcParams["figure.figsize"] = (20,7)
%matplotlib inline

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1

Write a Python function that calculates the roots of a quadratic equation $ax^2+bx+
c = 0$ of unknown $x$, using the quadratic formula below. 

$$x = \frac{-b\pm \sqrt{b^2-4ac}}{2a}$$

The function output should be a two element tuple. The first element, `elem_1`, should be the root computed by setting the $\pm$ to $+$, and the second element, `elem_2`, should be the root computed by setting the $\pm$ to $-$.

The discriminant is $b^2 - 4ac$. You may assume $b^2 - 4ac \ge 0$.

We will cover tuples next week. For now, all you need to know is when you have multiple outputs from a function, Python returns all the values as a sequence of values called a tuple.

_Type your answer here, replacing this text._

```python
# ANSWER CELL
def solve_quadratic(a, b, c):
    """
    Arguments
    ---------
    a : coefficient of x**2
    b : coefficient of x
    c : constant
    
    Returns
    -------
    elem_1 : larger root of quadratic equation
    elem_2 : smaller root of quadratic equation
    """
    discriminant = ...
    elem_1 = ...
    elem_2 = ...
    return elem_1, elem_2 
```

```python
solve_quadratic(6,-17,12)
```

```python
grader.check("q1")
```

## Question 2.0

Carbon dating is a method for determining the age of an object containing the radioisotope carbon-14. Since carbon-14 is subject to radioactive decay, the number of carbon-14 atoms in a sample will decay over time according to the following exponential decay law:

$$ N(t) = N_0 e^{-\lambda t} $$

where $N_0$ is the number of atoms of the isotope in the original sample, $N(t)$ is the remaining number of atoms after a time $t$, $\lambda$ is the exponential decay constant and $t$ is the time.

Write a function that calculates the remaining fraction of carbon-14 in a sample after a certain amount of time (in years). The remaining fraction is defined as the ratio of the number of atoms at the present time ($N(t)$) to the initial number of atoms ($N_0$). You can re-arrange the decay formula above to find an expression for this ratio.

Assume $\lambda = 0.00012097 \: \text{years}^{-1}$.

```python
# ANSWER CELL
def c14_dating(time):
    ratio = ...
    return ratio
```

Try out your new `c14_dating()` function in the cell below for `time = 1` year. Feel free to try it for other inputs.

You do not need to store any outputs from the cell below. This is just to allow you to test your function.

```python

```

```python
grader.check("q2.0")
```

## Question 2.1

Create a simpler version of Q2.0 by writing a lambda function called `lambda_c14_dating()`. Your lambda function should perform the same calculation as Q2.0 with `time` as the input and fraction remaining as the output. Keep the decay constant as $\lambda = 0.00012097 \: \text{years}^{-1}$

```python
# ANSWER CELL
lambda_c14_dating = ...
```

Try out your new `lambda_c14_dating()` function in the cell below for `time = 1` year. (Just as you did for Q2.0.) Feel free to try it for other inputs.

You do not need to store any outputs from the cell below. This is just to allow you to test your function.

```python

```

```python
grader.check("q2.1")
```

## Question 2.2

Now create a more general version of the function in Q2.0 and name the new function `exp_decay()`. This new function will perform the same calculation as `c14_dating()` but will include two input arguments: `time` for $t$ and `decay_constant` for $\lambda$, the exponential decay constant. Call the argument for $\lambda$ `decay_constant`. Set the default value for `decay_constant` equal to the exponential decay constant for c-14, which is 0.00012097.

```python
# ANSWER CELL
...
```

Try out your new `exp_decay()` function in the cell below for `time = 1` year. Feel free to try it for other inputs.

You do not need to store any outputs from the cell below. This is just to allow you to test your function.

```python

```

```python
grader.check("q2.2")
```

## Question 2.3

Evaluate the fraction of carbon-14 remaining over the inclusive range of time (years) = [0, 100, 200, .., 9800, 9900, 10000].

Assign your time range to `time_array`. Note that your array should have 100 year increments and include both 0 and 10,000 years. Refer to lecture notes for further help.

Assign the values of the fraction of carbon-14 remaining to `c14frac_array`. Use your `exp_decay()` function from Q2.2 to evaluate `time_array` to calculate the fraction of carbon-14 remaining throughout the time range.

```python
# ANSWER CELL
time_array = ...
c14frac_array = ...
```

Print out your arrays in the cell below to see how the values change.

```python

```

```python
grader.check("q2.3")
```

## Question 3

A ball is launched in the air with initial velocity $v_o$ from an initial position of coordinates $(x_0, y_0)$ as illustrated below: 

<center> <img src="resources/E7_Lab2_1.jpg" style="width: 300px;"/> </center>
<center> Figure 1: Ball launching at angle $\theta$ from the x-axis with inital velocity $v_0$ </center>

The ball is launched at an angle $\theta$ (expressed in degrees) from the x-axis. The
equations describing the motion (i.e. the position of the ball $(x(t), y(t)$) as a function
of time $t$) of the ball are:

$$ x(t) = x_0 + v_0 t \cos(\theta) $$
$$ y(t) = y_0 + v_0 t \sin(\theta) - \frac{1}{2}gt^2 $$

where $g$ is the gravitational acceleration ($g = 9.81\,\text{m}/\text{s}^2$}). Since $\theta$ is given in degrees,
you will want to convert it first into radians before using the sine or cosine functions.

You may find this function helpful: `np.deg2rad()`, which takes as input an angle in degrees and converts it to radians.

**Question 3.0:** Write a function named `proj_time()` that calculates the time at which the ball reaches the ground. Note that when the ball hits the ground, $y = 0$. So, to find the time at which the ball reaches the ground you have to solve the quadratic equation: 
$$ - \frac{1}{2}gt^2 + v_0 t \sin(\theta)  + y_0 = 0 $$ 
You should have already defined a function that solves a quadratic equation (function `solve_quadratic()`, refer to Question 1). So, your new function `proj_time()` should have the following input arguments (use the same argument names):
* `y0`, which is the initial $y$ position
* `v0`, which is the initial velocity
* `theta`, which is the launching angle in degrees

Inside `proj_time()`, call the function `solve_quadratic()` with the correct inputs based on the quadratic equation above.  Recall that `solve_quadratic()` returns two elements. The correct solution for the time at which the ball reaches the ground is the second element, `elem_2`, for which the root is computed by setting the $\pm$ to $-$. Your function `proj_time()` should only return one solution, the correct one.

```python
# run this cell to create a global variable g, but make sure you don't have any other global variables also named g
g = 9.81
```

```python
# ANSWER CELL
...
```

Try out your new `proj_time()` function in the cell below. Try using `y0 = 0` meters, `v0 = 15` m/s, and `theta = 40` degrees. Feel free to try it for other input values.

You do not need to store any outputs from the cell below. This is just to allow you to test your function.

```python

```

```python
grader.check("q3.0")
```

**Question 3.1:**  Write a function named `proj_distance()` that calculates the horizontal distance traveled by the ball from its initial horizontal position, $x_0$, before it reaches the ground.Your function should take as input arguments `y0`, `v0`, and `theta`. You can first solve for the time when the ball hits the ground using the function above, `proj_time()`, and then use this time in the following equation to get the distance traveled: $$ x(t) = x_0 + v_0 t \cos(\theta) $$

```python
# ANSWER CELL
...
```

Try out your new `proj_distance()` function in the cell below. Try using `y0 = 0` meters, `v0 = 15` m/s, and `theta = 40` degrees. Feel free to try it for other input values.

You do not need to store any outputs from the cell below. This is just to allow you to test your function.

```python

```

```python
grader.check("q3.1")
```

**Question 3.2:** Calculate the horizontal distance traveled by a ball starting at $5\,\text{m}$ above the ground, with a velocity of $1\,\text{m}/\text{s}$ at an angle of 30 degrees above horizontal. Store the value as `q3_2`.

```python
# ANSWER CELL
q3_2 = ...
q3_2
```

```python
grader.check("q3.2")
```

**Question 3.3:** Reassign the gravitational acceleration variable, `g`, to $32.2\,\text{ft}/\text{s}^2$. Then calculate the horizontal distance traveled by a ball starting at $5\,\text{ft}$ above the ground, with a velocity of $1\,\text{ft}/\text{s}$ at an angle of 30 degrees above horizontal. Store the value as `q3_3`.

```python
# ANSWER CELL
...
```

```python
# Do not remove. Run this cell to reassign g to 9.81 m/s^2.
g = 9.81
```

```python
grader.check("q3.3")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing lab 2!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 4)
```
>>> assert get_hash(solve_quadratic(5, -23, 2)) == '582fa9d143bfec5b7bfba577d5839f89'
>>> assert get_hash(solve_quadratic(1,2,1)) == 'bb6dedb4a3fcd11f0b5570754ca0fc80'
>>> assert get_hash(np.round(solve_quadratic(5.2,3.5,-2.1),5)) == '6bf522b49df28612d661b9e4dca578a3'
```


Tests for q2.0 (points: 3)
```
>>> assert get_hash(round(c14_dating(5730),4)) == 'd310cb367d993fb6fb584b198a2fd72c'
>>> assert get_hash(round(c14_dating(0.956),4)) == '94eab0ef89d927c724320fec6169d935'
>>> assert get_hash(round(c14_dating(14023),4)) == 'adc61c91420dc18855e1b8f869d62f02'
```


Tests for q2.1 (points: 3)
```
>>> assert get_hash(round(lambda_c14_dating(5730),4)) == 'd310cb367d993fb6fb584b198a2fd72c'
>>> assert get_hash(round(lambda_c14_dating(0.956),4)) == '94eab0ef89d927c724320fec6169d935'
>>> assert get_hash(round(lambda_c14_dating(14023),4)) == 'adc61c91420dc18855e1b8f869d62f02'
```


Tests for q2.2 (points: 3)
```
>>> assert get_hash(round(exp_decay(10000,decay_constant=0.005),4)) == '30565a8911a6bb487e3745c0ea3c8224'
>>> assert get_hash(round(exp_decay(5700,decay_constant=0.0005),4)) == 'f7c0bea0f6b8b0afb21183bd1fba47b5'
>>> assert get_hash(round(exp_decay(6000),4)) == '65f9014d65dbcfbfa19c0f7fa6ab4da8' , 'check defult argument'
>>> assert get_hash(round(exp_decay(1,decay_constant=0.5),4)) == '8a52647f06947d3d26632ec454dd25d7'
>>> assert get_hash(round(exp_decay(106.34,decay_constant=0.0034),4)) == '58374b5008429c8708b89e0ff3cef80a'
>>> assert get_hash(round(exp_decay(1456),4)) == 'ae9975b561135ecc09b35e7d203f41b3'
```


Tests for q2.3 (points: 3)
```
>>> assert get_hash(time_array[0]) == 'cfcd208495d565ef66e7dff9f98764da' , 'Check start value of time_array'
>>> assert get_hash(time_array[-1]) == 'b7a782741f667201b54880c925faec4b' , 'Check final value of time_array'
>>> assert get_hash(np.sum(time_array)) == '39d9d6852c0ddbf22d973f6d3c72deae' , 'Check steps within time_array'
```
```
>>> assert get_hash(round(c14frac_array[1],4)) =='ccf8e6e10af7c853137169f67a523087'
>>> assert get_hash(round(c14frac_array[-1],4)) == 'e152a2147de532f22e5acd4d4b252c84'
>>> assert get_hash(round(np.sum(c14frac_array),4)) == 'f82762b52fbe1dde5c5a5100ad7d2775'
```


Tests for q3.0 (points: 3)
```
>>> g = 9.81
>>> assert get_hash(proj_time(y0=0, v0=15, theta=40)) == '085a4d0b759822dd9066e0f3eebe96b5'
>>> assert get_hash(proj_time(y0=10, v0=0, theta=45)) == '5d54d3cd4e1e02079de389a99a64931f'
```


Tests for q3.1 (points: 3)
```
>>> g = 9.81
>>> assert get_hash(proj_distance(y0=0, v0=15, theta=40)) == '901c8912f3bd90f4c881badd02f571cb'
>>> assert get_hash(proj_distance(y0=10,v0=0, theta=45)) == '30565a8911a6bb487e3745c0ea3c8224'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(np.round(q3_2,4)) == 'c5204578f71b466f2818d519b95a925e'
```


Tests for q3.3 (points: 2)
```
>>> assert get_hash(np.round(q3_3,4)) == 'd65764cf7615df6f1cc880fdfb53c669'
```
