```python
# Initialize Otter
import otter
grader = otter.Notebook("lab06.ipynb")
```

# E7: Lab Assignment 06

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to familiarize you with plotting as well as reading and writing data.

## Autograder
You will be provided with some test cases (i.e., sample input data) and associated answers (i.e., expected outputs) that you can use to help check your code. The provided test cases are not exhaustive, and it is your responsibility to ensure that your code works in general, not just for a few supplied test cases. We may use additional hidden test cases in grading your lab assignments.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Your answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
Question| Points
--- | --
1 | 5.0
2 | 7.0
3 | 8.0
4 | 5.0
5 | 0.0
Total | 25.0

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import math
import numpy as np
import matplotlib.pyplot as plt
import cartopy.crs as ccrs
np.seterr(all='ignore');

import hashlib
def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Heated Copper Bar

In this problem you will examine the heat distribution in a one-dimensional copper bar over time. In particular, you will create a visual representation of the temperature $u(x,t)$ at different points $x$ along the bar as time $t$ progresses. The copper bar shown in Figure 1 has a fixed length $l$ and known initial temperature distribution $u(x,0) = f(x)$. The two ends are maintained at a constant temperature of 0&deg;C.

<img src='resources/rod.png' width=55%/>
 <figcaption style="text-align:center"><strong> Figure 1.  1D bar with homogeneous boundary conditions,</strong> <a href="https://en.wikipedia.org/wiki/Heat_equation#/media/File:Temp_Rod_homobc.svg">Wikipedia</a></figcaption>   

The bar has the following physical properties:
* length $l = 100\,\mathrm{cm}$
* constant mass density $\rho = 8.92 \,\mathrm{g/cm^3}$
* specific heat capacity $c  = 0.092 \,\mathrm{cal/( g\,^\circ C)}$
* thermal conductivity $k = 0.95 \, \mathrm{cal/(cm\,s\,^{\circ}C})$

The thermal diffusivity $\alpha$ is related to these physical properties through the equation $$\alpha = \dfrac{k}{c \rho}$$
Assume the initial temperature distribution is
$$ u(x,0) = f(x) = 100 \sin\left(\frac{3\pi x}{l}\right),  \label{eq:u0}\tag{1} $$
where $x$ is the horizontal coordinate along the bar. 
The temperature distribution $u(x,t)$ at point $x$ and time $t$ under these conditions is then 

$$ u(x,t) =  \exp\left(-\lambda^2 t\right)u(x,0),  \label{eq:uxt}\tag{2} $$
where we've defined
$$\lambda^2 = \alpha\left(\dfrac{3\pi}{l}\right)^2$$

### Question 1.0

Assign the bar length to the variable `length` and define a lambda function `u0(x)` which takes in:
* `x`: `float` or `np.ndarray`
     * coordinates along the bar, $x \in [0,l]$

and returns the intital temperature distribution shown in Eq.$\eqref{eq:u0}$ at each point in `x` for this particular problem.

Once you are done, test your `u0` function for `x = np.linspace(0,length)` and assign it to the variable `q1_0`. Feel free to test other inputs.

```python
# ANSWER CELL

# Bar length, cm
length = ...

# Initial temperature distribution function u0(x)
u0 = ...
```

```python
# TEST YOUR FUNCTION HERE
q1_0 = ...
```

```python
grader.check("q1.0")
```

### Question 1.1

Assign the bar density, specific heat and thermal conductivity to the variables `rho`,`c` and `k`, respectively. Similarly, assign $\alpha$ to `alpha` and $\lambda^2$ to `lam_sq` using the definitions in the problem statement. Lastly, define a lambda function `u(x,t)` which takes in:
* `x`: `float` or `np.ndarray` 
    * coordinates along the bar, $x \in (0,l)$ 
* `t`:  `float` 
    * time,  $ t > 0$ 
    
and returns the temperature distribution at time `t` shown in Eq.2 at each point in `x` for this particular problem.

Once you are done, test your `u(x,t)` function for `x = np.linspace(0,length)` and `t = 10` and assign it to the variable `q1_1`.  Feel free to test other inputs.

```python
# ANSWER CELL

# Density, g/cm3
rho = ...

# Specific heat, cal/g/deg C
c = ...

# Thermal conductivity, cal/cm/s/deg C
k = ...

# Thermal diffusivity
alpha = ...

# Lambda squared
lam_sq = ...

# Temperature distribution function  u(x,t)
u = ...
```

```python
# TEST YOUR FUNCTION HERE
q1_1 = ...
```

```python
grader.check("q1.1")
```

### Question 1.2

The next step is to compute the temperature on a discrete one-dimensional grid of coordinates $x$ at different times $t$. Using your `u(x,t)` function, define the following arrays:
*  `x` :  grid of 1000 equally spaced points in the range $[0,l]$ 
* `u_t0` : temperature $u(x,t)$ on the grid `x` at time $t = 0\,\mathrm{s}$  
* `u_t50` : temperature $u(x,t)$ on the grid `x` at time $t = 50\,\mathrm{s}$  
* `u_t150` :  temperature $u(x,t)$ on the grid `x` at time $t = 150\,\mathrm{s}$

```python
# Define x grid 
x = ...

# u(x,t) at t = 0s
u_t0 = ...

# u(x,t) at t = 50s
u_t50 = ...

# u(x,t) at t = 150s
u_t150 = ...
```

```python
# Do not modify this cell for grading purposes
x_q1 = np.copy(x)
```

```python
grader.check("q1.2")
```

### Question 1.3

We are now ready to plot the temperature distributions. Create a `matplotlib.pyplot` figure stored in `fig_1` and perform the following tasks on the same figure:
* Plot `u_t0` vs. `x` with a blue (`b`) solid (`-`) line . 
* Plot `u_t50` vs. `x` with a red (`r`) dashed (`--`) line. 
* Plot `u_t150` vs. `x` with a green (`g`) dash-dotted (`-.`) line. 
* Set the figure title to  "Temperature distribution along copper bar" with font size 16.
* Set the x-axis label to "x, cm" with font size 14.
* Set the y-axis label to "Temperature, degrees Celsius" with font size 14
* Set the x-axis limits to (0,100).
* Set the y-axis limits to (-120,120).
* Add a legend to the figure. The labels for the three line plots should respectively be "t = 0s", "t = 50s" and "t = 100s".
* Add grid lines to the figure.

When you are finished, your figure should look like Figure 2. Feel free to experiment with plotting options that have not been explicitly specified.

<img src='resources/q1.png' width = 60%/>
<figcaption style="text-align:center"><strong> Figure 2. Question 1 plot.</strong></figcaption>   

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Create figure 
fig_1 = ...

# Plot u(x,t) vs. x for different t values
...
...
...

# Set figure title
...

# Set axes labels
...
...

# Set axis limits
...
...

# Add legend
...

# Add gridlines
...

plt.show()
```

```python
grader.check("q1.3")
```

## Question 2: Vibrating Membrane 
In this problem you will examine the natural modes of vibration of a rectangular membrane of length $a$ and width $b$, fixed along all four edges. Vibrating plates or membranes are commonly encountered in various fields such as structural engineering, mechanical engineering and acoustics. Understanding their mode shapes is crucial for analyzing their behavior and resonance frequencies.

Assuming the membrane obeys the two-dimenional wave equation, the (normalized) amplitude $f$ of the $(m,n)$ natural mode of vibration is

$$f\left(x,y,a,b\,; m,n\right) = \sin\left(\frac{m \pi}{a}x\right)\sin\left(\frac{n \pi}{b}y\right),\label{eq:modeshape}\tag{3}$$

where $x$ and $y$ are coordinates along the membrane and $m$ and $n$ are positive integer mode indices.

### Question 2.0

Write a function `mode_shape(x,y,a,b,m,n)` which takes in:
* `x`: `float` or `np.ndarray`
     * $x$ coordinates along the membrane, $x \in (0,a)$
* `y`: `float` or `np.ndarray`
     * $y$ coordinates along the membrane, $y \in (0,b)$
* `a`: `float`
     * membrane length, $a > 0$
* `b`: `float`
     * membrane width, $b > 0$
* `m`: `int`
     * mode index along x, $m = 1,2,3,\dots$
* `n`: `int`
     * mode index along y, $n = 1,2,3,\dots$

and returns the mode amplitude $f\left(x,y,a,b\,; m,n\right)$ shown in Eq.$\eqref{eq:modeshape}$ at each point defined by `x` and `y`.


Once you are done, test your `mode_shape` function for `x = np.linspace(0,1)`,`y = np.linspace(0,1)`,`a = 1`,`b = 1`,`m = 1`,`n = 1` and assign it to the variable `q2_0`.

```python
# ANSWER CELL
```

```python
# TEST YOUR FUNCTION HERE
q2_0 = ...
```

```python
grader.check("q2.0")
```

### Question 2.1

Let us now compute the $m = n = 1$ natural mode of vibration for a membrane of length $a = 3$ and width $b = 2$ on a discrete two-dimensional grid. To this end, first create two linear ranges of 1000 equally spaced points each for the $x$ and $y$ coordinates, named `x` and `y`, respectively. Then, define combined two-dimensional coordinate arrays `X` and `Y` for each point on the grid. The documentation of [`np.meshgrid`](https://numpy.org/doc/stable/reference/generated/numpy.meshgrid.html) may be helpful.

Finally, using your `mode_shape` function, compute the mode amplitude on the `X,Y` grid and assign the output array to the variable `Z`.

```python
# ANSWER CELL

# Length
a = ...
# Width
b = ...
# Parameter m
m = ...
# Parameter n
n = ...

# Create x- and y-ranges
x = ...
y = ...

# Create combined X and Y grid
X,Y = ...

# Compute Z
Z = mode_shape(X,Y,a,b,m,n)
```

```python
# Do not modify and run this cell for grading purposes
x_q2_1 = np.copy(x)
y_q2_1 = np.copy(y)
X_q2_1 = np.copy(X)
Y_q2_1 = np.copy(Y)
Z_q2_1 = np.copy(Z)
```

```python
grader.check("q2.1")
```

### Question 2.2

Visualize the membrane's natural mode of vibration computed in Question 2.1 by creating a surface plot of `Z` on the `X,Y` grid. The documentation of [`plt.plot_surface`](https://matplotlib.org/stable/api/_as_gen/mpl_toolkits.mplot3d.axes3d.Axes3D.plot_surface.html#mpl_toolkits.mplot3d.axes3d.Axes3D.plot_surface) may be helpful.

Create a `matplotlib.pyplot` figure stored in `fig_2` and perform the following tasks on the same figure:
* Create a surface plot of the natural mode shape for $a = m = 3$ and $b= n = 2$, i.e. `Z` vs. `X` and `Y`.
* Set the surface colormap to `plasma`.
* Set the figure title to  "Natural mode shape for m = 3, n = 2" with font size 14.
* Set the x-axis label to "X" with font size 12.
* Set the y-axis label to "Y" with font size 12.
* Set the y-axis label to "Z" with font size 12.
* Adjust the plot to have equal axis aspect ratios, refer to the [documentation.](https://matplotlib.org/stable/gallery/subplots_axes_and_figures/axis_equal_demo.html)
* Adjust the viewing distance if some of the labels are clipped from the figure. Take a look at the `zoom` argument of [set_box_axpect.](https://matplotlib.org/stable/api/_as_gen/mpl_toolkits.mplot3d.axes3d.Axes3D.set_box_aspect.html)

When you are finished, your figure should look like Figure 3. Feel free to experiment with plotting options that have not been explicitly specified.

<img src="resources/q2_2.png" width=45%/>
<figcaption style="text-align:center"><strong> Figure 3. Question 2.2 plot.</strong></figcaption>   

```python
# ANSWER CELL

# Create figure
fig_2 = ...

# Create axes with 3d projection
ax = ...

# Plot surface
...

# Set title
...

# Set labels
...
...
...

# Set axes aspect ratio
...

plt.show() 
```

```python
grader.check("q2.2")
```

### Question 2.3

Create another set of two-dimensional coordinate grids `X` and `Y` of 1000 equally spaced points in each direction for a square membrane of unit length $a = b = 1$. Then, create a `matplotlib.pyplot` figure stored in `fig_3` and perform the following tasks on the same figure:
* Create a 2x2 subplot layout and plot the membrane mode shape for $m ,n = {1,2}$. Store the subplot axes in the variable `axs`. The subplot in the $m$-th row and $n$-th column should contain the mode shape for the respective $m$ and $n$ values. For example, the subplot in the lower left (row 2, column 1) should contain the plot for $m = 2$ and $n =1$.
* Set the figure title (suptitle) to "Mode shapes" with font size 16.
* For each plot do the following:
    * Set the surface colormap to `coolwarm`
    * Set the subplot title "m = {m}, n = {n}" where you insert the appropriate value in the braces. See Figure 4 for reference.
    * Set the x-axis label to "X".
    * Set the y-axis label to "Y".
    * Set the y-axis label to "Z".
    * Adjust the viewing distance if some of the labels are clipped from the subplot. Take a look at the `zoom` argument of [set_box_axpect.](https://matplotlib.org/stable/api/_as_gen/mpl_toolkits.mplot3d.axes3d.Axes3D.set_box_aspect.html)


When you are finished, your figure should look like Figure 4. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q2_3.png" width = 60%/>
<figcaption style="text-align:center"><strong> Figure 4. Question 2.3 plot.</strong></figcaption>   
</figure>

```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Create x and y 1D grids
x = ...
y = ...

# Create combined x and y 2D grid
X,Y = ...

# Create figure of 2x2 3D subplots
fig_3, axs = ...

# Plot mode shape on each subplot

# Set figure title
...

plt.show()
```

```python
# Do not modify and run this cell for grading purposes
x_q2_3 = np.copy(x)
y_q2_3 = np.copy(y)
X_q2_3 = np.copy(X)
Y_q2_3 = np.copy(Y)
```

```python
grader.check("q2.3")
```

## Question 3: Wind over North America

In this problem you will examine and visualize a sample dataset of air currents over North America. 

### Question 3.0

The `resources` folder contains the following CSV data files on a 35 x 41 grid:

* `wind_x.csv`: $x$ coordinates (longitude) 
* `wind_y.csv`: $y$ coordinates (latitude) 
* `wind_u.csv`: $x$ components of velocity
* `wind_v.csv`: $y$ components of velocity

First, read the data files listed above using [`np.loadtxt`](https://numpy.org/doc/stable/reference/generated/numpy.loadtxt.html) and assign the arrays to the variables `wind_x`, `wind_y`, `wind_u` and `wind_v`, respectively. Make sure to specify the `delimiter` argument which is the character used to separate the values in the file. Recall that CSV stands for "comma-separated values".

*Hint:* The relative file path to `wind_x.csv`, for example, is `resources/wind_x.csv`.

```python
# ANSWER CELL

# Read wind_x data
wind_x = ...

# # Read wind_y data
wind_y = ...

# # Read wind_u data
wind_u = ...

# # Read wind_v data
wind_v = ...
```

```python
grader.check("q3.0")
```

### Question 3.1

Compute the magnitude of the wind velocity at each grid point and assign it to the variable `wind_mag`. Make sure to preserve the shape of the wind velocity arrays. The magnitude of a 2D vector $\vec{v}$ with components $v_x$ and $v_y$ is defined as:

$$ ||\vec{v}|| = \sqrt{v_x^2 + v_y^2}$$

```python
# ANSWER CELL

# Compute magnitude of wind velocity on the grid
wind_mag = ...
```

```python
grader.check("q3.1")
```

### Question 3.2

Write the array `wind_mag` to `resources/wind_mag.csv` using [`np.savetxt`](https://numpy.org/doc/stable/reference/generated/numpy.savetxt.html). Use a comma (`,`) delimiter.

```python
# Write the wind magnitude data to a csv file
...
```

```python
grader.check("q3.2")
```

### Question 3.3

Now, we'll plot the wind data overlaid on a map of North America with the help of the `cartopy` module (refer to [Chapter 12.3](https://pythonnumericalmethods.berkeley.edu/notebooks/chapter12.03-Working-with-Maps.html) of the course textbook).

After you've familiarized yourself with `cartopy`, create a `matplotlib.pyplot` figure stored in `fig_4` and perform the following tasks on the same figure:
* Create figure axes using the Plate Carrée projection and draw coastlines on the plot. 
* Create a filled contour plot (`contourf`) of the velocty magnitude and assign it to the variable `wind_contour`. Note that the data is already provided on a 2D meshgrid. Set the number of levels to 20.
* Add a colorbar of the contour plot to the figure and assign it to the variable `cbar_wind`. Set the scale (`shrink`) factor to an appropriate value.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to "Wind velocity magnitude".
* Set the x and y axes limits/extent to the minimum and maximum values in `wind_x` and `wind_y`, respectively. Make sure to specify the `crs` (coordinate reference system) argument of the `set_extent` method to be the `ccrs.PlateCarree()` projection like the axes. This way the coordinates will be aligned with the figure regardless of its dimensions.
 
When you are finished, your figure should look like Figure 5. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q3_3.png" width = 60%/>

<figcaption style="text-align:center"><strong>Figure 5. Question 3.3 plot.</strong></figcaption>   
 

```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_4 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Create wind magnitude filled contour plot
wind_contour = ...

# Create colorbar
cbar_wind = ...

# Set limits/extent of plot
...

# Add grid lines
grid = ...

# Set title
...

plt.show()
```

```python
grader.check("q3.3")
```

### Question 3.4

Lastly, we'll plot the wind vector field using `quiver`. Create a `matplotlib.pyplot` figure stored in `fig_5` and perform the following tasks on the same figure:
* Create axes using the Plate Carrée projection and draw coastlines on the plot.
* Create a vector plot (`quiver`) of the velocty field and assign it to the varaible `vel_quiver`.  Note that the data is already provided on a 2D meshgrid. Feel free to experiment with the colormap but use the `wind_mag` array as the colormapping values of the arrows. 
* Add a colorbar of the contour plot to the figure and assign it to the variable `cbar_wind`. Set the title to "Magnitude" and the scale (`shrink`) factor to an appropriate value.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to "Wind velocity field".
* Set the x and y axes limits/extent to the minimum and maximum values in `wind_x` and `wind_y`, respectively. Make sure to specify the `crs` (coordinate reference system) argument of the `set_extent` method to be the `ccrs.PlateCarree()` projection like the axes. This way the coordinates will be aligned with the figure regardless of its dimensions.
 
When you are finished, your figure should look like Figure 6. Feel free to experiment with plotting options that have not been explicitly specified.

<img src="resources/q3_4.png" width = 60%/>

<figcaption style="text-align:center"><strong>Figure 6. Question 3.4 plot.</strong></figcaption>   


```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_5 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Create wind magnitude filled contour plot
vel_quiver = ...

# Create colorbar
cbar_wind = ...

# Set limits/extent of plot
...

# Add grid lines
grid = ...

# Set title
...

plt.show()
```

```python
grader.check("q3.4")
```

## Problem 4: Earthquakes

In this problem you will analyze and visualize earthquake data for the continental United States in 2022.

### Question 4.0

The `resources/earthquakes_2022.csv` file contains records of earthquakes with magnitude 3.5 and above in the continental United States during 2022. The data file has 229 rows, each representing a different earthquake event, with the following data in 3 columns :

1. decimal degrees latitude
2. decimal degrees longitude
3. the magnitude of the event

The columns are ordered as listed above. Read the data file using [`np.loadtxt`](https://numpy.org/doc/stable/reference/generated/numpy.loadtxt.html) and assign the array to the variable `eq_data`. Make sure to specify the `delimiter` argument which is the character used to separate the values in the file. Recall that CSV stands for "comma-separated values".

```python
# ANSWER CELL

# Read data
eq_data = ...

# Display data
eq_data
```

```python
grader.check("q4.0")
```

### Question 4.1

Create a `matplotlib.pyplot` figure stored in `fig_6` and perform the following tasks on the same figure:
* Create a histogram using [`plt.hist`](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.hist.html) with 5 bins in the range (3,8) and assign the output to the variables `values, bins, bars`. Set the bin edge color to black (`k`).
* Set the title to "Continental US earthquakes in 2022".
* Set the x-axis label to "Magnitude".
* Set the y-axis label to "Frequency".
* Label the histogram bars with the count in each bin (*Hint:* Refer to the documentation of `bar_label`.).
 
When you are finished, your figure should look like Figure 10. Feel free to experiment with plotting options that have not been explicitly specified.

<img src="resources/q4_1.png" width = 50%/>

<figcaption style="text-align:center"><strong>Figure 7. Question 4.1 plot.</strong></figcaption>   


```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt

# Create figure
fig_6 = ...

# Plot histogram
values, bins, bars = ...
 
# Set title and labels
...
...
...

# Add bar labels
...

plt.show()
```

```python
grader.check("q4.1")
```

### Question 4.2

We will now zoom in on the state of California and the surrounding area. Create a `matplotlib.pyplot` figure stored in `fig_7` and perform the following tasks on the same figure:
* Create axes using the Plate Carrée projection and draw coastlines on the plot.
* Create a scatter plot (`scatter`) of the earthquake location and assign it to the variable `eq_scatter`. 
* Use the magnitude data for the color mapping and set the transparency (`alpha`) of the markers to 0.5. 
* Use the magnitude data to the fourth power (`**4`) for the marker size mapping.
* Add a colorbar to the figure and assign it to the variable `cbar_eq`. Set the scale (`shrink`) factor to an appropriate value and set the label to 'Magnitude'.
* Add a grid to the figure with transparency (`alpha`) of 0.5 and assign it to the variable `grid`. The gird lines should also be dotted (`:`) and black (`k`). Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the documentation and [examples](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html).  Also, draw longitude and latitude tick labels only on the bottom and left sides. Refer to the [documentation](https://scitools.org.uk/cartopy/docs/latest/gallery/gridlines_and_labels/gridliner.html). 
* Set the title to 'CA area earthquakes in 2022'.
* To create a rough bounding box around California:
    * set the x-axis limits to (-126,-114).
    * set the y-axis limits to (32,42).
    * Make sure to specify the `crs` (coordinate reference system) argument of the `set_extent` method to be the `ccrs.PlateCarree()` projection like the axes. This way the coordinates will be aligned with the figure regardless of its dimensions.
 
When you are finished, your figure should look like Figure 8. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q4_2.png" width = 60%/>

<figcaption style="text-align:center"><strong>Figure 8. Question 4.2 plot.</strong></figcaption>   


```python
# ANSWER CELL

# Do not modify these lines for grading purposes
import matplotlib.pyplot as plt
import cartopy.crs as ccrs

# Create figure
fig_7 = ...

# Create axes using Plate Carrée projection
ax = ...

# Add coastilines
...

# Scatter plot of earthquakes
eq_scatter = ...

# Create colorbar
cbar_eq = ...

# Set axes extent/limits    
...

# Add grid lines
grid = ...

# Set title
...

# Plot UC Berkeley on the map, no need to modify this code
berkeley_lon, berkeley_lat = -122.2585, 37.8719
berkeley_blue = np.array([0, 50, 98])/255
berkeley_gold = np.array([253, 181, 21])/255
ax.scatter(berkeley_lon, berkeley_lat, marker='*', ec=berkeley_blue, 
           c=[berkeley_gold], linewidth=1.4, s=400, zorder=2)
ax.text(berkeley_lon - 0.3, berkeley_lat - 0.3, 'UC Berkeley',
         horizontalalignment='right',color = berkeley_blue, weight="bold")

plt.show()
```

```python
grader.check("q4.2")
```

## (Optional) Question 5: Potential flow around a circular cylinder
 
Potential flow around a circular cylinder is a classical problem in fluid dynamics. It inovles the flow of an inviscid, incompressible fluid around a circular cylinder transverse to the flow. Far from the cylinder, the flow is unidirectional and uniform with a constant velocity $U$ in the x-direction. The flow has no vorticity and thus the velocity field is irrotational which allows for closed-form analytical solutions for quantities of interest.

*Note:* This question is not graded and is provided solely for additional practice.

<img src="resources/flow.gif" width=50%/>
 <figcaption style="text-align:center"><strong> Figure 9. Potential flow around a circular cylinder,</strong> <a href="https://en.wikipedia.org/wiki/Potential_flow_around_a_circular_cylinder#/media/File:Inviscid_flow_around_a_cylinder.gif">Wikipedia</a></figcaption>   

### Question 5.0

First, we'll determine the pressure field. Let the coordinate system be centered on the cicular cylinder with radius $R$. The solution for the (gauge) *pressure* field $p$ in polar coordinates assuming unit density is then 
$$ p = \frac{1}{2} U^2 \left( 2 \frac{R^2}{r^2}\cos\left(2\theta\right) - \frac{R^4}{r^4}\right) \label{eq:pressure}\tag{4} $$ 
Another quantity of interest is the so-called *stream function* $\psi$. It can be used to plot streamlines, which represent the trajectories of particles in a steady flow. The stream function for this problem is
$$\psi  = U\left(r-{\frac {R^{2}}{r}}\right)\sin \theta \label{eq:stream}\tag{5} $$
Finally, we are also interested in the *velocity potential* $\phi$. We'll use it to plot equipotential lines along which the fluid has uniform velocity. The velocity potential for this problem is
$$ \phi = Ur\left(1+{\frac {R^{2}}{r^{2}}}\right)\cos \theta  \label{eq:velpotential}\tag{6} $$

Let the cylinder radius have radius $R = 1$ and the far field fluid velocity be $U = 10$. You are tasked with computing the three quantities described above on a uniform two-dimensional **polar** grid. The polar coordinates are the radius $r$ and angle $\theta$. The grid should have 1000 equally spaced points in the radial direction within the range $\left[R, 10R\right]$ and 1000 equally spaced points in the angular direction within the range $\left[0, 2\pi\right]$.

```python
# Fluid velocity away from the cylinder
U = ...

# Cylinder radius
R = ...

# Create r and theta 1D grids
theta = ...
rr = ...

# Create meshgrid in two dimensions
...

# Compute pressure around the cylinder, Equation 4
pressure = ...

# Compute stream function, Equation 5
stream = ...

# Compute velocity potential , Equation 6
velocity_potential = ...
```

```python
grader.check("q5.0")
```

### Question 5.1

Now, we will plot the pressure, streamlines and velocty equipotential lines. First, we'll need to convert the 2D polar grid to Cartesian coordinates. The relatonship between polar and Cartesian coordiantes is
\begin{align}
x &= r \cos{\theta} \\
y &= r \sin{\theta}
\end{align}

After you convert the grid, create a `matplotlib.pyplot` figure stored in `fig_8` and perform the following tasks on the same figure:
* Create a filled contour plot (`contourf`) of the pressure field and assign it to the variable `pressure_contour`. Set the number of levels to 50 and the colormap to `jet`.
* Add a colorbar of the pressure field plot to the figure and assign it to the variable `cbar_pressure`. The colorbar should have only two ticks - at the minimum and maximum value which occur in the plot and the labels should respectively be "Low" and "High".
* Create a contour plot (`contour`) of the streamlines and assign it to the variable `stream_contour`.. Set the number of levels to 50, the line width to 0.8 and the line color to black (`k`).
* Create a contour plot (`contour`) of the velocity potential and assign it to the variable `velocity_contour`. Set the number of levels to 50, the line width to 0.8 and the line color to white (`w`).
* Set the title to "Pressure field" with font size 14.
* Set the x-axis label to "x" with font size 12.
* Set the y-axis label to "y" with font size 12.
* Set the x-axis limits to (-5,5).
* Set the y-axis limits to (-5,5).
* Remove the ticks from the x and y axes.

*Hint:* The `negative_linestyle` may be useful for the `contour` plots.


When you are finished, your figure should look like Figure 10. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q5_1.png" width = 50%/>
 <figcaption style="text-align:center"><strong> Figure 10. Question 5.1 plot.</strong></figcaption>   


```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Convert polar to cartesian coordinates
X, Y = ...

# Create figure
fig_8 = ...

# Plot circle/cylinder, do not modify
th =  np.linspace(0,2*np.pi,100)  
xunit = np.cos(th)
yunit = np.sin(th)
h = plt.plot(xunit, yunit,'-k')

# Pressure plot
pressure_contour = ...

# Create colorbar and set the ticks and tick labels
cbar_pressure = ...
...

# Streamline contour plot
stream_contour = ...

# Velocity equipotential lines contour plot
velocity_contour = ...

# Set title
...

# Set axes labels
...
...

# Set axes ticks
...
...

# Set axes limits 
...
...

plt.show()
```

```python
# Do not modify and run this cell for grading purposes
X_q5_1 = np.copy(X)
Y_q5_1 = np.copy(Y)
```

```python
grader.check("q5.1")
```

### Question 5.2

Finally, we will plot the fluid velocity magnitude and vector field around the cylinder. The velocty vector components in polar coordiantes are

$$v_{r}={\frac {\partial \phi }{\partial r}}=U\left(1-{\frac {R^{2}}{r^{2}}}\right)\cos \theta, \label{eq:vr}\tag{7}$$

$$v_{\theta }={\frac {1}{r}}{\frac {\partial \phi }{\partial \theta }}=-U\left(1+{\frac {R^{2}}{r^{2}}}\right)\sin \theta, \label{eq:vt}\tag{8} $$

Compute these quantities on the 2D polar grid from Question 3.0 and do not modify the code which converts them to the Cartesian components $v_x$ and $v_y$. The magnitude of a 2D vector $\vec{v}$ with components $v_x$ and $v_y$ is defined as:
$$ ||\vec{v}|| = \sqrt{v_x^2 + v_y^2} , \label{eq:vnorm}\tag{9}$$

When you are done, create a `matplotlib.pyplot` figure stored in `fig_9` and perform the following tasks on the same figure:
* Create a filled contour plot (`contourf`) of the velocty magnitude $\eqref{eq:vnorm}$ and assign it to the variable `v_contour`. Set the number of levels to 30 and the colormap to `jet`.
* Add a colorbar of the velocity magniude plot to the figure and assign it to the variable `cbar_vel`. The colorbar should have only two ticks - at the minimum and maximum value which occur in the plot and the labels should respectively be "Low" and "High".
* Create a vector plot (`quiver`) of the velocity field $v_x, v_y$ and assign it to the variable `v_quiver`. Becasue the grid is so dense, plot at only every 25th grid point in each direction (*Hint:* Use step slicing). Set the scale factor to 200.
* Set the title to "Velocity field" with font size 14.
* Set the x-axis label to "x" with font size 12.
* Set the y-axis label to "y" with font size 12.
* Set the x-axis limits to (-3,3).
* Set the y-axis limits to (-3,3).
* Remove the ticks from the x and y axes.

*Hint:* The `negative_linestyle` may be useful for the `contour` plots.


When you are finished, your figure should look like Figure 11. Feel free to experiment with plotting options that have not been explicitly specified.


<img src="resources/q5_2.png" width = 50%/>
 <figcaption style="text-align:center"><strong> Figure 11. Question 5.2 plot.</strong></figcaption>   


```python
# ANSWER CELL

# Do not modify this line for grading purposes
import matplotlib.pyplot as plt

# Velocity in polar coordinates
vr = ...
vt = ...

# Velocity in cartesian coordinates, do not modify these lines
vx = vr*np.cos(t) - vt*np.sin(t)
vy = vr*np.sin(t) + vt*np.cos(t)

# Create figure
fig_9 = ...

# Plot circle/cylinder, do not modify
th =  np.linspace(0,2*np.pi,100)  
xunit = np.cos(th)
yunit = np.sin(th)
h = plt.plot(xunit, yunit,'-k')

# Compute velocity magnitude
v_mag = ...

# Velocity magnitude filled contour plot
v_contour = ...
                  
# Add velocity magnitude colorbar
cbar_vel = ...
...
                  
# Velocity vector field quiver plot
step = ...
v_quiver = ...

# Set title
...

# Set axes labels
...
...

# Set axes ticks
...
...

# Set axes limits 
...
...

plt.show()
```

```python
grader.check("q5.2")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
import matplotlib.image as mpimg
img = mpimg.imread('resources/animal.png')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congrats on finishing Lab 6!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 1)
```
>>> assert get_hash(length) == 'f899139df5e1059396431415e770c6dd', 'Check the bar length.'
```
```
>>> assert get_hash(type(q1_0)) != '14e736438b115821cbb9b7ac0ba79034', 'Make sure to test your function!'
```
```
>>> assert get_hash(np.round(u0(0), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your u0 function.'
>>> assert get_hash(np.round(u0(10), decimals=2)) == 'bf9cba947baf9cc43f447e0e0fc7efce', 'Check your u0 function.'
>>> assert get_hash(np.round(u0(60), decimals=2)) == '9f7c34c89841d32c138e886829eb4154', 'Check your u0 function.'
>>> assert get_hash(np.round(u0(100), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your u0 function.'
```
```
>>> x = np.linspace(0, 100)
>>> assert get_hash(len(u0(x))) == 'c0c7c76d30bd3dcaefc96f40275bdc0a', 'Check your u0 function.'
>>> assert get_hash(np.round(np.sum(u0(x)), decimals=2)) == '82fd5432008dcbb742f92940ad26cd05', 'Check your u0 function.'
```


Tests for q1.1 (points: 1)
```
>>> assert get_hash(rho) == '1637533ef837128ee0199716d73687bd', 'Check rho.'
>>> assert get_hash(k) == '248a7444f08189bb31ba143eabebe4e5', 'Check k.'
>>> assert get_hash(c) == '052fbf0fbad7b4a752d34e1dc0e76cfa', 'Check c.'
>>> assert get_hash(np.round(alpha, decimals=2)) == '695ec0ee7352c033530309d0ad3418cb', 'Check alpha.'
>>> assert get_hash(np.round(lam_sq, decimals=2)) == '04817efd11c15364a6ec239780038862', 'Check lam_sq.'
```
```
>>> assert get_hash(type(q1_1)) != '14e736438b115821cbb9b7ac0ba79034', 'Make sure to test your function!'
```
```
>>> assert get_hash(np.round(u(0, 100), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your u function.'
>>> assert get_hash(np.round(u(10, 10), decimals=2)) == '19ca1e8274ab2a38604c8baa84ef7baa', 'Check your u function.'
>>> assert get_hash(np.round(u(60, 30), decimals=2)) == '89dbc31f785433f9aeefa440b26f01a7', 'Check your u function.'
>>> assert get_hash(np.round(u(100, 100), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your u function.'
```
```
>>> x = np.linspace(0, 100)
>>> assert get_hash(len(u(x, 10))) == 'c0c7c76d30bd3dcaefc96f40275bdc0a', 'Check your u function.'
>>> assert get_hash(np.round(np.sum(u(x, 10)), decimals=2)) == 'dd665bc9865f92f5c247d793c1f5c019', 'Check your u function.'
```


Tests for q1.2 (points: 1)
```
>>> assert get_hash(len(x_q1)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check the x grid.'
>>> assert get_hash(np.round(np.sum(x_q1), decimals=2)) == '3b68e4b87a954d6116e95334c6c59068', 'Check the x grid.'
>>> assert get_hash(np.round(np.sum(u_t0), decimals=2)) == '47bff7aaf1ec63d5d71200ec4177fd02', 'Check u_t0.'
>>> assert get_hash(np.round(np.sum(u_t50), decimals=2)) == '643c90867bbecb6e12ceb9c61a9fc180', 'Check u_t50.'
>>> assert get_hash(np.round(np.sum(u_t150), decimals=2)) == '9f985f523aa3f1bc1e9c410ccd67f1fe', 'Check u_t150.'
```


Tests for q1.3 (points: 2)
```
>>> assert fig_1.axes[0].get_lines()[0].get_linestyle() == '-', 'Check the line styles.'
>>> assert fig_1.axes[0].get_lines()[1].get_linestyle() == '--', 'Check the line styles.'
>>> assert fig_1.axes[0].get_lines()[2].get_linestyle() == '-.', 'Check the line styles.'
>>> assert fig_1.axes[0].get_lines()[0].get_color() == 'b', 'Check the line colors.'
>>> assert fig_1.axes[0].get_lines()[1].get_color() == 'r', 'Check the line colors.'
>>> assert fig_1.axes[0].get_lines()[2].get_color() == 'g', 'Check the line colors.'
>>> assert fig_1.axes[0].get_lines()[0].get_label() == 't = 0s', 'Check the line labels.'
>>> assert fig_1.axes[0].get_lines()[1].get_label() == 't = 50s', 'Check the line labels.'
>>> assert fig_1.axes[0].get_lines()[2].get_label() == 't = 150s', 'Check the line labels.'
```
```
>>> assert 'X' in fig_1.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert 'CM' in fig_1.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert fig_1.axes[0].xaxis.label.get_fontsize() == 14, 'Check the x-axis label fontsize.'
>>> assert 'TEMPERATURE' in fig_1.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert 'DEGREES' in fig_1.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert 'CELSIUS' in fig_1.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert fig_1.axes[0].get_ylabel() == 'Temperature, degrees Celsius', 'Check the y-axis label.'
>>> assert fig_1.axes[0].yaxis.label.get_fontsize() == 14, 'Check the y-axis label fontsize.'
>>> assert 'TEMPERATURE' in fig_1.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'DISTRIBUTION' in fig_1.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'COPPER' in fig_1.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'BAR' in fig_1.axes[0].get_title().upper(), 'Check the title label.'
>>> assert fig_1.axes[0].title.get_fontsize() == 16, 'Check the title label fontsize.'
```
```
>>> assert type(fig_1.axes[0].get_legend()) != type(None), 'Make sure to add a legend.'
>>> assert fig_1.axes[0].get_xlim() == (0.0, 100.0), 'Check the x-axis limits.'
>>> assert fig_1.axes[0].get_ylim() == (-120.0, 120.0), 'Check the y-axis limits.'
```


Tests for q2.0 (points: 1)
```
>>> assert get_hash(type(q2_0)) != '14e736438b115821cbb9b7ac0ba79034', 'Make sure to test your function!'
```
```
>>> assert get_hash(np.round(mode_shape(0.2, 0.6, 1, 1, 1, 1), decimals=2)) == '2363c78ab7dba59b8443d958b47cfa2b', 'Check your mode_shape function.'
>>> assert get_hash(np.round(mode_shape(0.23, 0.7, 1, 1, 3, 4), decimals=2)) == 'ab01272978f4106a12ede6fe2a850be4', 'Check your mode_shape function.'
>>> assert get_hash(np.round(mode_shape(0.3, 0.3, 1, 1, 8, 1), decimals=2)) == '671a647819140975b0649cb3d2f4804a', 'Check your mode_shape function.'
>>> assert get_hash(np.round(mode_shape(0.7, 0.8, 1, 1, 5, 5), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your mode_shape function.'
```
```
>>> x = np.linspace(0, 100)
>>> assert get_hash(len(mode_shape(x, x, 1, 1, 1, 1))) == 'c0c7c76d30bd3dcaefc96f40275bdc0a', 'Make sure your mode_shape function works for x and y arrays.'
>>> assert get_hash(np.round(np.sum(mode_shape(x, x, 1, 1, 1, 1)), decimals=2)) == '42b4dad1f1f82501da84d8942aa3343b', 'Make sure your mode_shape function works for x and y arrays.'
```


Tests for q2.1 (points: 1)
```
>>> assert get_hash(len(x_q2_1)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check your 1D x grid.'
>>> assert get_hash(np.round(np.sum(x_q2_1), decimals=2)) == '1c51ac3926a046148a24ba048afc85cd', 'Check your 1D x grid.'
>>> assert get_hash(len(y_q2_1)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check your 1D y grid.'
>>> assert get_hash(np.round(np.sum(y_q2_1), decimals=2)) == '1434f81012ce9c2aba692d5177889ef4', 'Check your 1D y grid.'
>>> assert get_hash(X_q2_1.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D x grid.'
>>> assert get_hash(np.round(np.sum(X_q2_1), decimals=2)) == 'ead4a206965a49fc387fd5b4de90867c', 'Check your 2D X grid.'
>>> assert get_hash(Y_q2_1.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D y grid.'
>>> assert get_hash(np.round(np.sum(Y_q2_1), decimals=2)) == 'be6a5089b8e4ba482c427df57552ad5f', 'Check your 2D Y grid.'
>>> assert get_hash(Z_q2_1.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D y grid.'
>>> assert get_hash(np.round(np.sum(Z_q2_1), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your 2D Z grid.'
```


Tests for q2.2 (points: 2)
```
>>> assert fig_2.axes[0].get_children()[0].get_cmap().name == 'plasma', 'Check colormap.'
>>> assert fig_2.axes[0].get_aspect() == 'equal', 'Check the axes aspect ratio.'
>>> assert 'X' in fig_2.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert fig_2.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> assert 'Y' in fig_2.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert fig_2.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> assert 'Z' in fig_2.axes[0].get_zlabel().upper(), 'Check the z-axis label.'
>>> assert fig_2.axes[0].zaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> assert 'NATURAL' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'MODE' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'SHAPE' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'M' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert '3' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'N' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert '2' in fig_2.axes[0].get_title().upper(), 'Check the title label.'
>>> assert fig_2.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
```


Tests for q2.3 (points: 3)
```
>>> assert get_hash(len(x_q2_3)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check your 1D x grid.'
>>> assert get_hash(np.round(np.sum(x_q2_3), decimals=2)) == '02541b7cbb840b5f6bd8eff3c9f56f61', 'Check your 1D x grid.'
>>> assert get_hash(len(y_q2_3)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check your 1D y grid.'
>>> assert get_hash(np.round(np.sum(y_q2_3), decimals=2)) == '02541b7cbb840b5f6bd8eff3c9f56f61', 'Check your 1D y grid.'
>>> assert get_hash(X_q2_3.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D x grid.'
>>> assert get_hash(np.round(np.sum(X_q2_3), decimals=2)) == '21819d66f20373895e5f413cfe0e5c0f', 'Check your 2D x grid.'
>>> assert get_hash(len(Y_q2_3)) == 'a9b7ba70783b617e9998dc4dd82eb3c5', 'Check your 2D y grid.'
>>> assert get_hash(np.round(np.sum(Y_q2_3), decimals=2)) == '21819d66f20373895e5f413cfe0e5c0f', 'Check your 2D y grid.'
```
```
>>> assert fig_3._suptitle.get_text() == 'Mode shapes', 'Check figure suptitle.'
>>> assert fig_3._suptitle.get_fontsize() == 16, 'Check figure suptitle font size.'
>>> assert axs.shape == (2, 2), 'Make sure you have a 2x2 subplot layout stored in the variable "axs".'
>>> lst = [1, 2]
>>> for m in lst:
...     for n in lst:
...         assert axs[m - 1, n - 1].get_xlabel().upper() == 'X', 'Check x-axis labels.'
...         assert axs[m - 1, n - 1].get_ylabel().upper() == 'Y', 'Check y-axis labels.'
...         assert axs[m - 1, n - 1].get_zlabel().upper() == 'Z', 'Check z-axis labels.'
...         assert axs[m - 1, n - 1].get_title() == 'm = {m}, n = {n}'.format(m=m, n=n), 'Check title labels.'
...         assert axs[m - 1, n - 1].get_children()[0].get_cmap().name == 'coolwarm', 'Check colormap.'
```


Tests for q3.0 (points: 1)
```
>>> assert isinstance(wind_x, np.ndarray), 'Make sure wind_x is assigned and is a numpy array.'
>>> assert isinstance(wind_y, np.ndarray), 'Make sure wind_y is assigned and is a numpy array.'
>>> assert isinstance(wind_u, np.ndarray), 'Make sure wind_u is assigned and is a numpy array.'
>>> assert isinstance(wind_v, np.ndarray), 'Make sure wind_v is assigned and is a numpy array.'
```
```
>>> assert get_hash(np.round(np.sum(wind_x), decimals=2)) == '191f944a56e709fed4cbba9cbe4c8bd0', 'Check the wind_x data.'
>>> assert get_hash(np.round(np.sum(wind_y), decimals=2)) == 'cf87726119a3db49884cc2e364dbe32a', 'Check the wind_y data.'
>>> assert get_hash(np.round(np.sum(wind_u), decimals=2)) == '02e71c8357f8983537872ad2f30580e4', 'Check the wind_u data.'
>>> assert get_hash(np.round(np.sum(wind_v), decimals=2)) == 'dda186858e2cec344497f604ebb95dab', 'Check the wind_v data.'
```


Tests for q3.1 (points: 1)
```
>>> assert get_hash(np.round(np.sum(wind_mag), decimals=2)) == '872e27abbbc3735010e8d1f58ae6c465', 'Check the wind magnitude computation.'
```


Tests for q3.2 (points: 1)
```
>>> q3_2 = np.loadtxt('resources/wind_mag.csv', delimiter=',')
>>> assert get_hash(np.round(np.sum(q3_2), decimals=2)) == '872e27abbbc3735010e8d1f58ae6c465', 'Check the write command.'
```


Tests for q3.3 (points: 3)
```
>>> assert wind_contour.get_cmap().name == 'coolwarm', 'Check wind_contour colormap.'
>>> assert 15 < len(wind_contour.levels) < 25, 'Check velocitycontour levels.'
```
```
>>> assert 'WIND' in fig_4.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'VELOCITY' in fig_4.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'MAGNITUDE' in fig_4.axes[0].get_title().upper(), 'Check the title label.'
>>> assert fig_4.axes[0].get_xlim() == (-134.3, -70.188), 'Check the x-axis limits.'
>>> assert fig_4.axes[0].get_ylim() == (17.5, 60.0), 'Check the y-axis limits.'
```
```
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5, 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':', 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k', 'Check grid color'
```


Tests for q3.4 (points: 2)
```
>>> assert np.isclose(np.sum(vel_quiver.X), np.sum(wind_x), rtol=0.01), 'Check quiver plot data.'
>>> assert np.isclose(np.sum(vel_quiver.Y), np.sum(wind_y), rtol=0.01), 'Check quiver plot data.'
>>> assert np.isclose(np.sum(vel_quiver.U), np.sum(wind_u), rtol=0.01), 'Check quiver plot data.'
>>> assert np.isclose(np.sum(vel_quiver.V), np.sum(wind_v), rtol=0.01), 'Check quiver plot data.'
```
```
>>> assert 'WIND' in fig_5.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'VELOCITY' in fig_5.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'FIELD' in fig_5.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'MAGNITUDE' in cbar_wind.ax.get_ylabel().upper(), 'Check the colorbar label.'
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5, 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':', 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k', 'Check grid color'
>>> assert fig_5.axes[0].get_xlim() == (-134.3, -70.188), 'Check the x-axis limits. Make sure to specify "crs".'
>>> assert fig_5.axes[0].get_ylim() == (17.5, 60.0), 'Check the y-axis limits. Make sure to specify "crs".'
```
```
>>> import matplotlib
>>> assert np.isclose(np.sum(matplotlib.cm.ScalarMappable(norm=vel_quiver.norm, cmap=vel_quiver.cmap).to_rgba(wind_mag).flatten()), np.sum(vel_quiver.get_facecolor().flatten()), rtol=0.01), 'Use wind_mag as the colormap values.'
```


Tests for q4.0 (points: 1)
```
>>> assert get_hash(eq_data.shape) == '86475183e983421e2c569d248a56dd25', 'Check your read command.'
```


Tests for q4.1 (points: 1)
```
>>> assert get_hash(len(values)) == 'e4da3b7fbbce2345d7772b0674a318d5', 'Check bin count.'
>>> assert get_hash(len(bins)) == '1679091c5a880faf6fb5e6087eb1b2dc', 'Check bin count.'
>>> assert get_hash(sum(values)) == '7359bd4fcd36e2be25916898a1e291b0', 'Check histogram data.'
>>> assert get_hash(sum(bins)) == '12e9099d922d8928c4087e78fffc4975', 'Check bin range.'
```
```
>>> assert 'MAGNITUDE' in fig_6.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert 'FREQUENCY' in fig_6.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert 'CONTINENTAL' in fig_6.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'EARTHQUAKES' in fig_6.axes[0].get_title().upper(), 'Check the title label.'
>>> assert '2022' in fig_6.axes[0].get_title().upper(), 'Check the title label.'
```


Tests for q4.2 (points: 3)
```
>>> assert np.isclose(np.sum(eq_scatter.get_offsets()[:, 0]), np.sum(eq_data[:, 1]), rtol=0.01), 'Check scatter plot data.'
>>> assert np.isclose(np.sum(eq_scatter.get_offsets()[:, 1]), np.sum(eq_data[:, 0]), rtol=0.01), 'Check scatter plot data.'
```
```
>>> assert np.isclose(np.sum(eq_scatter.get_sizes()), np.sum(eq_data[:, 2] ** 4), rtol=0.01), 'Check the marker size mapping.'
>>> import matplotlib
>>> assert np.isclose(np.sum(matplotlib.cm.ScalarMappable(norm=eq_scatter.norm, cmap=eq_scatter.cmap).to_rgba(eq_data[:, 2], alpha=0.5).flatten()), np.sum(eq_scatter.get_facecolor().flatten()), rtol=0.01), 'Use the earthquake magnitude as the colormap values. Make sure you set the color alpha as well.'
```
```
>>> assert 'CA' in fig_7.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'EARTHQUAKES' in fig_7.axes[0].get_title().upper(), 'Check the title label.'
>>> assert '2022' in fig_7.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'MAGNITUDE' in cbar_eq.ax.get_ylabel().upper(), 'Check the colorbar label.'
>>> assert grid.top_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.bottom_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.left_labels == True, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.right_labels == False, 'Make sure to have labels only on the bottom and left sides.'
>>> assert grid.collection_kwargs['alpha'] == 0.5, 'Check grid alpha'
>>> assert grid.collection_kwargs['linestyle'] == ':', 'Check grid linestyle'
>>> assert grid.collection_kwargs['color'] == 'k', 'Check grid color'
>>> assert np.isclose(np.sum(fig_7.axes[0].get_xlim()), np.sum([-126, -114])), 'Check the x-axis limits.'
>>> assert np.isclose(np.sum(fig_7.axes[0].get_ylim()), np.sum([32, 42])), 'Check the y-axis limits.'
```


Tests for q5.0 (points: 0)
```
>>> assert get_hash(np.round(np.sum(theta), decimals=2)) == '9311ad9f26394d518925abac95b99be8', 'Check theta 1D grid.'
>>> assert get_hash(np.round(np.sum(rr), decimals=2)) == '381a7a26ff2526446cbb97fee6f8505a', 'Check r 1D grid.'
>>> assert get_hash(np.round(np.sum(t), decimals=2)) == 'b9d0eb1983a1cf8354d55656a60c03fb', 'Check r 2D grid.'
>>> assert get_hash(np.round(np.sum(r), decimals=2)) == '35ce9af3f1187f980f3b1bebcd574a7b', 'Check theta 2D grid.'
```
```
>>> assert get_hash(np.round(np.sum(pressure), decimals=2)) == '81ba63f8e5af6c7eb551de20fc6f7088', 'Check pressure 2D grid.'
>>> assert get_hash(np.round(np.sum(stream), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check stream 2D grid.'
>>> assert get_hash(np.round(np.sum(velocity_potential), decimals=2)) == '55aa3a294d30e95cb206ecfc1b85dee1', 'Check velocity_potential 2D grid.'
```


Tests for q5.1 (points: 0)
```
>>> assert get_hash(X_q5_1.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D x grid.'
>>> assert get_hash(np.round(np.sum(X_q5_1), decimals=2)) == '381a7a26ff2526446cbb97fee6f8505a', 'Check your 2D x grid.'
>>> assert get_hash(Y_q5_1.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D y grid.'
>>> assert get_hash(np.round(np.sum(Y_q5_1), decimals=2)) == '30565a8911a6bb487e3745c0ea3c8224', 'Check your 2D y grid.'
```
```
>>> assert len(cbar_pressure.get_ticks()) == 2, 'Check colorbar ticks.'
>>> assert np.isclose(cbar_pressure.get_ticks()[0], -150, rtol=0.01), 'Check colorbar ticks.'
>>> assert np.isclose(cbar_pressure.get_ticks()[1], 50, rtol=0.01), 'Check colorbar ticks.'
>>> assert cbar_pressure.ax.get_yticklabels()[0].get_text().upper() == 'LOW', 'Check colorbar ticks.'
>>> assert cbar_pressure.ax.get_yticklabels()[1].get_text().upper() == 'HIGH', 'Check colorbar ticks.'
>>> assert pressure_contour.get_cmap().name == 'jet', 'Check pressure_contour colormap.'
>>> assert 45 < len(pressure_contour.levels) < 55, 'Check pressure_contour levels.'
>>> assert stream_contour.linewidths == 0.8, 'Check stream_contour line width.'
>>> assert 45 < len(stream_contour.levels) < 55, 'Check stream_contour levels.'
>>> assert stream_contour.tcolors[0][0] == (0, 0, 0, 1), 'Check stream_contour line colors.'
>>> assert velocity_contour.linewidths == 0.8, 'Check velocity_contour line width.'
>>> assert 45 < len(velocity_contour.levels) < 55, 'Check velocity contour levels.'
>>> assert velocity_contour.tcolors[0][0] == (1, 1, 1, 1), 'Check stream_contour line colors.'
```
```
>>> assert 'X' in fig_8.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert fig_8.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> assert 'Y' in fig_8.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert fig_8.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> assert 'PRESSURE' in fig_8.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'FIELD' in fig_8.axes[0].get_title().upper(), 'Check the title label.'
>>> assert fig_8.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
>>> assert fig_8.axes[0].get_xlim() == (-5, 5), 'Check the x-axis limits.'
>>> assert fig_8.axes[0].get_ylim() == (-5, 5), 'Check the y-axis limits.'
>>> assert len(fig_8.axes[0].get_xticks()) == 0, 'Check the x-axis ticks.'
>>> assert len(fig_8.axes[0].get_yticks()) == 0, 'Check the y-axis ticks.'
```


Tests for q5.2 (points: 0)
```
>>> assert get_hash(vr.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D vr grid.'
>>> assert get_hash(np.round(np.sum(vr), decimals=2)) == '75cd8970acd3ddaa0d10134c11c32ce7', 'Check your 2D vr grid.'
>>> assert get_hash(vt.shape) == '2c279ef3676bba689eabb103f0a319e8', 'Check your 2D vt grid.'
>>> assert get_hash(np.round(np.sum(vt), decimals=2)) == 'e80c77464a0cfaf2c46a5574f6bc76b9', 'Check your 2D vt grid.'
```
```
>>> assert v_quiver.scale == 200, 'Checl quiver plot scale.'
>>> assert len(v_quiver.X) == 1600, 'Make sure to plot at every 25th grid point only.'
>>> assert get_hash(np.round(np.sum(v_quiver.X), decimals=2)) == '5ab881cdc85e2fc2abe9a30734de1b5b', 'Make sure to plot at every 25th grid point only.'
>>> assert len(v_quiver.Y) == 1600, 'Make sure to plot at every 25th grid point only.'
>>> assert get_hash(np.round(np.sum(v_quiver.Y), decimals=2)) == 'bb42d52e328cb7ecb4a0732d34fd1a71', 'Make sure to plot at every 25th grid point only.'
```
```
>>> assert len(cbar_vel.get_ticks()) == 2, 'Check colorbar ticks.'
>>> assert np.isclose(cbar_vel.get_ticks()[0], 0, rtol=0.01), 'Check colorbar ticks.'
>>> assert np.isclose(cbar_vel.get_ticks()[1], 19.99, rtol=0.01), 'Check colorbar ticks.'
>>> assert cbar_vel.ax.get_yticklabels()[0].get_text() == 'Low', 'Check colorbar ticks.'
>>> assert cbar_vel.ax.get_yticklabels()[1].get_text() == 'High', 'Check colorbar ticks.'
>>> assert v_contour.get_cmap().name == 'jet', 'Check pressure_contour colormap.'
>>> assert 25 < len(v_contour.levels) < 35, 'Check pressure_contour levels.'
>>> assert stream_contour.linewidths == 0.8, 'Check stream_contour line width.'
>>> assert 45 < len(stream_contour.levels) < 55, 'Check stream_contour levels.'
>>> assert stream_contour.tcolors[0][0] == (0, 0, 0, 1), 'Check stream_contour line colors.'
>>> assert velocity_contour.linewidths == 0.8, 'Check velocity_contour line width.'
>>> assert 45 < len(velocity_contour.levels) < 55, 'Check velocitycontour levels.'
>>> assert velocity_contour.tcolors[0][0] == (1, 1, 1, 1), 'Check stream_contour line colors.'
```
```
>>> assert 'X' in fig_9.axes[0].get_xlabel().upper(), 'Check the x-axis label.'
>>> assert fig_9.axes[0].xaxis.label.get_fontsize() == 12, 'Check the x-axis label fontsize.'
>>> assert 'Y' in fig_9.axes[0].get_ylabel().upper(), 'Check the y-axis label.'
>>> assert fig_9.axes[0].yaxis.label.get_fontsize() == 12, 'Check the y-axis label fontsize.'
>>> assert 'VELOCITY' in fig_9.axes[0].get_title().upper(), 'Check the title label.'
>>> assert 'FIELD' in fig_9.axes[0].get_title().upper(), 'Check the title label.'
>>> assert fig_9.axes[0].title.get_fontsize() == 14, 'Check the title label fontsize.'
>>> assert fig_9.axes[0].get_xlim() == (-3, 3), 'Check the x-axis limits.'
>>> assert fig_9.axes[0].get_ylim() == (-3, 3), 'Check the y-axis limits.'
>>> assert len(fig_9.axes[0].get_xticks()) == 0, 'Check the x-axis ticks.'
>>> assert len(fig_9.axes[0].get_yticks()) == 0, 'Check the y-axis ticks.'
```
