```python
# Initialize Otter
import otter
grader = otter.Notebook("lab01.ipynb")
```

# E7: Lab Assignment 01

You must submit the lab to Gradescope by the due date. You will submit the zip file produced by running the final cell of the assignment.

## About this Lab
The objective of this assignment is to familiarize you with basic operations in Python as well as introduce you to NumPy arrays.

## Autograder
For each question, you will have to fill in one or more Python functions. We provide an autograder with a number of test cases that you can use to test your function. Note that the fact that your function works for all test cases thus provided does necessarily guarantee that it will work for all possible test cases relevant to the question. It is your responsibility to test your function thoroughly, to ensure that it will also work in situations not covered by the test cases provided.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. `...` is just a placeholder.

You answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

If you want to read documentation on a Python function, you can type `help()` or `help(function)`.

## Score Breakdown
|Question    | Points |
| :--------  | -----: |
| 1.0 – 1.8  | 9.0    |
| 2.0        | 0.5    |
| 2.1        | 1.5    |
| 2.2 – 2.4  | 3.0    |
| 3.0 – 3.5  | 6.0    |
| 4.0 – 4.4  | 5.0    |
| Total      | 25.0   |

```python
# Please run this cell, and do not modify the contents
import math
import numpy as np
import cmath
import re
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import ipywidgets as widgets
from ipywidgets import FileUpload
from IPython.display import display
from PIL import Image
import os
import resources
from resources.utils import *

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Using Python as a calculator

In this question, we will be exploring how we can use Python as a calculator. 

**Question 1.0:** Define variables $a = 3, b = -6, c = 4$, and $x = 2$




```python
# ANSWER CELL
a = ...
b = ...
c = ...
x = ...
print(a,b,c,x)
```

```python
grader.check("q1.0")
```

For each of the following mathematical expressions, construct a one-line Python expression that computes the value and assigns it to the provided variable. Your Python expressions should give appropriate values regardless of the values of $a, b, c$, and $x$ (i.e. you should use the variable names instead of their numerical values).

**Question 1.1:** $$\sqrt{a^2+b^2+c^2}$$

*Hint:* Take a look at `math`'s `sqrt` function [here](https://docs.python.org/3/library/math.html#power-and-logarithmic-functions).

```python
# ANSWER CELL
q1_1 = ...
q1_1
```

```python
grader.check("q1.1")
```

**Question 1.2:** $$\frac{-b-\sqrt{b^2-4ac}}{2a}$$

Note that `math.sqrt()` does not support complex numbers. If you try to use  `math.sqrt()`, it will raise a `ValueError`.

*Hint:* Take a look at `cmath`'s `sqrt` function that supports complex numbers [here](https://docs.python.org/3/library/cmath.html).

```python
# ANSWER CELL
q1_2 = ...
q1_2
```

```python
grader.check("q1.2")
```

**Question 1.3:** $$\log_{10} \left(3 \left|x\right|+\frac{c}{5}\right)$$

*Hint:* Take a look at `math`'s logarithmic documentation [here](https://docs.python.org/3/library/math.html#power-and-logarithmic-functions).

```python
# ANSWER CELL
q1_3 = ...
q1_3
```

```python
grader.check("q1.3")
```

**Question 1.4:** $$\left(ax+\frac{ab}{c}\right)^{1/3}$$

*Hint:* Take a look at `math`'s `pow` documentation [here](https://docs.python.org/3/library/math.html#power-and-logarithmic-functions) or use a built-in function.

```python
# ANSWER CELL
q1_4 = ...
q1_4
```

```python
grader.check("q1.4")
```

**Question 1.5:** $$\left(\cos\left(\frac{\sqrt{a}}{3}\pi\right)\right)^2+\cos\left(\left(\frac{\sqrt{a}}{3}\pi\right)^2\right)$$

*Hint:* Take a look at the `math` module's helpful functions and constants [here](https://docs.python.org/3/library/math.html).

```python
# ANSWER CELL
q1_5 = ...
q1_5
```

```python
grader.check("q1.5")
```

**Question 1.6:** $$i+e^{i\pi}$$

*Hint:* The `cmath` module has some helpful functions to deal with complex numbers. Check them out [here](https://docs.python.org/3/library/math.html).

*Hint:* In Python, put the letter `j` after a number to make the number complex.

```python
# ANSWER CELL
q1_6 = ...
q1_6
```

```python
grader.check("q1.6")
```

**Question 1.7:** $$\arcsin(\cos(\ln(x)))$$

*Hint:* Python naming conventions for functions may not match with typical handwritten conventions.

```python
# ANSWER CELL
q1_7 = ...
q1_7
```

```python
grader.check("q1.7")
```

**Question 1.8:** $$\frac{a+2c}{\sin\left(\dfrac{b+2c}{\sqrt{a^2+b^2+c^2}}\right)}$$

```python
# ANSWER CELL
q1_8 = ...
q1_8
```

```python
grader.check("q1.8")
```

## Question 2: Carbon Emissions Data Entry

In this question, you will look at the 2021 carbon dioxide emissions of China, Qatar, and the USA. The table below shows each country's population in millions of people and 2021 emissions of carbon dioxide in megatonnes of CO$_2$.

| Country | Population (millions) | CO$_2$ Emissions (MtCO$_2$) |
|---------|-----------------------|-----------------------------|
| China   | 1412                  | 11472                       |
| Qatar   |  2.7                  | 96                          |
| USA     |  332                  | 5007                        |

**Question 2.0:**  To store the population and emissions of each country as separate variables, which naming convention is most appropriate to name our variables?

Since you are comparing three countries, you will have a total of six variables. In the variable names, `pop` is an abbreviation of population.

A. pop-China, CO2-China, pop-Qatar, CO2-Qatar, pop-USA, CO2-USA \
B. China_pop , China_CO2 , Qatar_pop, Qatar_CO2, USA_pop, USA_CO2 \
C. China_pop , CO2_China , QatarPop , CO2emissionQatar, UnitedStates_pop, UnitedStates_CO2

Answer in the following cell. Your answer should be a string `"A"`, `"B"`, or `"C"`, not the full list of variable names. \
Remember to put quotes around your answer choice.

```python
# ANSWER CELL
q2_0 = ...
q2_0
```

```python
grader.check("q2.0")
```

**Question 2.1:** Create the six population and emission variables using the naming convention from q2_0 and the values from the table below. If the tabular value does not contain a decimal, do not include a decimal in the variable assignment.

Keep the given units when assigning the variables. For example, use 1412 for China's population, not 1,412,000,000.

| Country | Population (millions) | CO$_2$ Emissions (MtCO$_2$)|
|---------|-----------------------|---------------------------|
| China   | 1412                  | 11472                     |
| Qatar   |  2.7                  | 96                        |
| USA     | 332                   | 5007                      |


```python
# ANSWER CELL
...
```

```python
grader.check("q2.1")
```

**Question 2.2:** What is the data type of `Qatar_CO2`? Assign the data type to `q2_2`. Use a function to determine the data type. Do not just type the answer.

```python
# ANSWER CELL
q2_2 = ...
q2_2
```

```python
grader.check("q2.2")
```

**Question 2.3:** What is the value of `Qatar_pop` if it is converted to an integer? Assign the new integer value to `q2_3`. Do not overwrite the data type of `Qatar_pop` to integer; only `q2_3` should be an integer.

```python
# ANSWER CELL
q2_3 = ...
q2_3
```

```python
grader.check("q2.3")
```

**Question 2.4:** What is the value of `Qatar_CO2` if it is converted to a float? Assign the new float value to `q2_4`. Do not overwrite the data type of `Qatar_CO2` to float; only `q2_4` should be a float.

```python
# ANSWER CELL
q2_4 = ...
q2_4
```

```python
grader.check("q2.4")
```

## Question 3: Carbon Dioxide Analysis

Using the six variables defined in Question 2, **construct logical expressions** to answer the following true/false questions.

Your answers should be a logical expression, and the outputs should be boolean values.

**Question 3.0:** Calculate the CO2 emissions per capita for each country using the equation below and assign the values to the following variables: `USA_per_capita`, `China_per_capita`, `Qatar_per_capita`.

$$\text{per capita emissions } \left[\frac{\text{tonne CO$_2$}}{\text{person}}\right] = \frac{\text{CO$_2$ emissions [MtCO$_2$]}}{\text{population [millions]}}$$

Do not perform any unit conversions. A megatonne is 10$^6$ tonnes and cancels out with millions (10$^6$) of people.

```python
# ANSWER CELL
...
```

```python
grader.check("q3.0")
```

**Question 3.1:** Did the USA emit more carbon dioxide per capita than China in 2021?

Construct a logical expression to answer this question and assign it to `q3_1`.

```python
# ANSWER CELL
q3_1 = ...
q3_1
```

```python
grader.check("q3.1")
```

**Question 3.2:** Did the USA emit more carbon dioxide per capita than either China or Qatar emitted in 2021?

You should answer this by combining two logical expressions. The first expression, `q3_2a`, should compare China and the USA, and the second expression, `q3_2b`, should compare Qatar and the USA. The final answer should be stored in `q3_2c`.

```python
# ANSWER CELL
q3_2a = ...
q3_2b = ...
q3_2c = ...
q3_2c
```

```python
grader.check("q3.2")
```

**Question 3.3:** For comparison, in 2021, the global average emissions per capita value was 4.7 tonnes CO$_2$. Do China, Qatar, and the USA all have per capita emissions higher than the global average?

You should answer this by combining three logical expressions. The first expression, `q3_3a`, should compare China and the global average, the second expression, `q3_3b`, should compare Qatar and the global average, and the third expression `q3_3c`, should compare the USA and the global average. The final answer (do all three countries have above average emissions) should be stored in `q3_3d`.

```python
# ANSWER CELL
q3_3a = ...
q3_3b = ...
q3_3c = ...
q3_3d = ...
q3_3d
```

```python
grader.check("q3.3")
```

**Question 3.4:** In 2021, 7.5% of China's CO$_2$ emissions were due to cement production, while 1.7% of Qatar's CO$_2$ emissions were due to cement production. Do China and Qatar have equal per capita CO$_2$ emissions due to cement production?

Calculate China's per capita CO$_2$ emissions due to cement production in `q3_4a`. \
Calculate Qatar's per capita CO$_2$ emissions due to cement production in `q3_4b`. \
Provide the final answer in `q3_4c`.

```python
# ANSWER CELL
q3_4a = ...
q3_4b = ...
q3_4c = ...
q3_4c
```

```python
grader.check("q3.4")
```

**Question 3.5:** In 2021, 15% of China's CO$_2$ emissions were due to oil, while 45% of the USA's CO$_2$ emissions were due to oil. Does China have higher per capita CO$_2$ emissions due to oil than the USA?

Construct a logical expression to answer this question and assign it to `q3_5`.

```python
# ANSWER CELL
q3_5 = ...
q3_5
```

```python
grader.check("q3.5")
```

## Question 4: Carbon Emissions NumPy Arrays

In this question you will look again at the 2021 carbon dioxide emissions data but we will use NumPy arrays. Recall that the data set included the following (note that we have added data for Germany):

| Country | Population (millions) | CO$_2$ Emissions (MtCO$_2$) |
|:--------|:----------------------|:----------------------------|
| China   | 1412                  | 11472                       |
| Qatar   |  2.7                  | 96                          |
| USA     |  332                  | 5007                        |
| Germany |  3.4                  | 675                         |

**Question 4.0:** Instead of storing individual values to individual variables like we did in Question 2, a much more useful and efficient way to store data is using arrays. Specifically, we will use NumPy arrays.

Create two NumPy arrays, one for population (save it as `pop`) and one for emissions (save is as `CO2`). The values in the arrays should be in the same order and should have the same units as the table above. Refer to the lecture notes to see how to create NumPy arrays.

```python
# ANSWER CELL
pop = ...
CO2 = ...
print(pop)
print(CO2)
pop_q4_0 = pop.copy(); CO2_q4_0 = CO2.copy() # DO NOT MODIFY THIS LINE AND LEAVE IT AT THE END
```

```python
grader.check("q4.0")
```

**Question 4.1:** The correct population for Germany should be 83.4 million, not 3.4. Correct the last entry of the population data array to 83.4. Do not retype the array – use array indexing and assign the last element to the correct value.

```python
# ANSWER CELL
...
print(pop)
pop_q4_1 = pop.copy(); CO2_q4_1 = CO2.copy() # DO NOT MODIFY THIS LINE AND LEAVE IT AT THE END
```

```python
grader.check("q4.1")
```

**Question 4.2:** Let's append the data for another country, India to the population and emissions arrays. India has a population of 1408 million people and CO$_2$ emissions of 2710 MtCO$_2$. Using `np.append()`, append the population and emission values of India to `pop` and `CO2`, respectively. Assign the new arrays to `pop` and `CO2`.

```python
# ANSWER CELL
...
print(pop)
print(CO2)
```

```python
grader.check("q4.2")
```

**Question 4.3:** We previously defined the per capita emissions for each country as:

$$\text{per capita emissions } \left[\frac{\text{tonne CO$_2$}}{\text{person}}\right] = \frac{\text{CO$_2$ emissions [MtCO$_2$]}}{\text{population [millions]}}$$

Calculate the per capita emissions for each county by dividing the arrays `pop` and `CO2` and assign the output to a variable named `per_capita`.

```python
# ANSWER CELL
per_capita = ...
print(per_capita)
```

```python
grader.check("q4.3")
```

**Question 4.4:** Recall that the global average emissions per capita value in 2021 was 4.7 tonnes CO$_2$. Did each of the five countries have per capita emissions higher than the global average? 

You should answer this by creating a logical expression using the array `per_capita`. The output should be another array that includes `True` or `False` for each country, depending on whether each country had per capita emissions higher than the global average. This should be stored in `q4_4`.

```python
# ANSWER CELL
q4_4 = ...
q4_4
```

```python
grader.check("q4.4")
```

### You're done with this Lab!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp;<i class="fa fa-save" style="font-size:16px;"></i>&nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing your first lab!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the .zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 1)
```
>>> assert get_hash(a) == 'eccbc87e4b5ce2fe28308fd9f2a7baf3'
>>> assert get_hash(b) == '596a3d04481816330f07e4f97510c28f'
>>> assert get_hash(c) == 'a87ff679a2f3e71d9181a67b7542122c'
>>> assert get_hash(x) == 'c81e728d9d4c2f636f067f89cc14862c'
```


Tests for q1.1 (points: 1)
```
>>> assert get_hash(round(q1_1, 2)) == '10f8e19d19bf3772687b5457a689fa21'
```


Tests for q1.2 (points: 1)
```
>>> assert get_hash(q1_2) == 'eac1be4ceede350c62d93c8212e3e536'
```


Tests for q1.3 (points: 1)
```
>>> assert get_hash(round(q1_3, 3)) == '5e295664d633f88870fc93eda8cf42c9'
```


Tests for q1.4 (points: 1)
```
>>> assert get_hash(round(q1_4, 3)) == '14a742b8cccaf65017719898cf735d84'
```


Tests for q1.5 (points: 1)
```
>>> assert get_hash(round(q1_5, 4)) == '88491693334ee557168464e310349583'
```


Tests for q1.6 (points: 1)
```
>>> assert get_hash(round(np.real(q1_6),8)) == 'c54490d3480079138c8c027a87a366e3'
>>> assert get_hash(round(np.imag(q1_6),8)) == 'e4c2e8edac362acab7123654b9e73432'
```


Tests for q1.7 (points: 1)
```
>>> assert get_hash(round(q1_7,4)) == '3205c5a0e0c890e0113e87dd3b5d7349'
```


Tests for q1.8 (points: 1)
```
>>> assert (get_hash(round(q1_8, 3)) == 'c0857c1de732e77e7f99af3055dda037')
```


Tests for q2.0 (points: 0.5)
```
>>> assert len(q2_0.upper()) == 1 , 'Make sure you have just the letter of your answer choice in quotes'
>>> assert get_hash(q2_0.upper()) == '9d5ed678fe57bcca610140957afab571'
```


Tests for q2.1 (points: 1.5)
```
>>> var_names = ['China_pop' , 'China_CO2' , 'Qatar_pop' , 'Qatar_CO2' , 'USA_pop' , 'USA_CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(China_pop) == '0e4e946668cf2afc4299b462b812caca'
>>> assert get_hash(China_CO2) == 'b7852f3d3a775b5188139a2c07024397'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584'
>>> assert get_hash(USA_pop) == 'c042f4db68f23406c6cecf84a7ebb0fe'
>>> assert get_hash(USA_CO2) == '351869bde8b9d6ad1e3090bd173f600d'
```


Tests for q2.2 (points: 1)
```
>>> assert get_hash(type(q2_2)) == '71f335aebe2d722b3351e4c9710a7236' , 'Check which function you used.'
>>> assert get_hash(q2_2) == '2203e3c2293eb8c9bda5bcd97acbf884' , 'Does the value have a fractional component?'
```


Tests for q2.3 (points: 1)
```
>>> assert get_hash(q2_3) == 'c81e728d9d4c2f636f067f89cc14862c'
>>> assert get_hash(Qatar_pop) == '23cb05d4a2e5bf28363909d92cd244e2' , 'Do not reassign the variable'
```


Tests for q2.4 (points: 1)
```
>>> assert get_hash(q2_4) == '2993d6c487b8101c0701534ba325cb91'
>>> assert get_hash(Qatar_CO2) == '26657d5ff9020d2abefe558796b99584' , 'Do not reassign the variable'
```


Tests for q3.0 (points: 1)
```
>>> q3_0_vars = ['China_per_capita' , 'Qatar_per_capita' , 'USA_per_capita']
>>> assert all([var_name in globals() for var_name in q3_0_vars]), 'Check variable names' # Checks that the variables exist
>>> assert get_hash(China_per_capita) == '2e08fabaf262a0fea4da5c818c9894e2' , 'Check variable value'
>>> assert get_hash(Qatar_per_capita) == '2f56f1561f0c21041643671d4fd9f85a' , 'Check variable value'
>>> assert get_hash(USA_per_capita) == 'c89b4e5c27b4928beff41e7782689f91' , 'Check variable value'
```


Tests for q3.1 (points: 1)
```
>>> assert get_hash(q3_1) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.2 (points: 1)
```
>>> assert get_hash(q3_2a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_2b) == 'f8320b26d30ab433c5a54546d21f414c'
>>> assert get_hash(q3_2c) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.3 (points: 1)
```
>>> assert get_hash(q3_3a) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3b) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3c) == 'f827cf462f62848df37c5e1e94a4da74'
>>> assert get_hash(q3_3d) == 'f827cf462f62848df37c5e1e94a4da74'
```


Tests for q3.4 (points: 1)
```
>>> assert get_hash(round(q3_4a,3)) == '302ddfd6d63bbe5693b96981c129d3e2'
>>> assert get_hash(round(q3_4b,3)) == '70e3b658922334c673fa2959dbfdfbd6'
>>> assert get_hash(q3_4c) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q3.5 (points: 1)
```
>>> assert get_hash(q3_5) == 'f8320b26d30ab433c5a54546d21f414c'
```


Tests for q4.0 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_0) == '5e784975e5eb0ff3a28a2aad3953dcbb'
>>> assert get_hash(CO2_q4_0) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.1 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop_q4_1) == 'd0ecac9d532783c816b4995f1327a2e6'
>>> assert get_hash(CO2_q4_1) == '8f484d873ab271148f32f7d0921afb08'
```


Tests for q4.2 (points: 1)
```
>>> var_names = ['pop' , 'CO2']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(pop) == '881602830957a3cd082095af4601bb6e'
>>> assert get_hash(CO2) == 'e88a56268abff0794d40bf0f030ed684'
```


Tests for q4.3 (points: 1)
```
>>> var_names = ['per_capita']
>>> assert all([var_name in globals() for var_name in var_names]), 'Check variable names'
>>> assert get_hash(per_capita) == '7b6ee17515b452692f05c2c2341c91cf'
```


Tests for q4.4 (points: 1)
```
>>> assert get_hash(q4_4) == '977c8f03a4a2c1813c6bcaaa0b831963'
```
