```python
# Initialize Otter
import otter
grader = otter.Notebook("lab00.ipynb")
```

# E7: Lab Assignment 00

Welcome to E7: Introduction to Computer Programming for Scientists and Engineers! You must submit the lab to Gradescope by the due date. You will submit the **zip file** produced by running the final cell of the assignment.

## About Lab 00
Each week you will complete a lab assignment like this one. You cannot learn to program just by watching someone else do it. You will need to be actively engaged in the learning process to develop your own understanding of the material through coding practice in lab sections and assignments. This is just a very simple, one-question assignment. The objective of this assignment is to acquaint you with the process of working on lab assignments and familiarize you with the procedure for submitting them.

## Autograder
You will be provided with some test cases (i.e., sample input data) and associated answers (i.e., expected outputs) that you can use to help check your code. The provided test cases are not exhaustive, and it is your responsibility to ensure that your code works in general, not just for a few supplied test cases. We may use additional hidden test cases in grading your lab assignments.

**Run the first cell, Initialize Otter**, to import the autograder and submission exporter.

## Answer cells
Throughout the assignment, replace `...` with your answers. We use `...` as a placeholder and theses should be deleted and replaced with your answers.

Your answers must be in the cells marked `# ANSWER CELL`, including your final calculation/code. However, do not perform scratchwork in `# ANSWER CELL`. Add a new cell to perform your scratchwork. Your scratchwork will not be graded and does not need to be included in your submission unless otherwise noted.

To read the documentation on a Python function, you can type `help()` and add the function name between parentheses.

## Score Breakdown
|Question    | Points |
| :--------  | -----: |
| 1.0        | 5.0    |
| Total      | 5.0    |

**Run the cell below**, to import the required modules.

```python
# Please run this cell, and do not modify the contents
import re
import matplotlib.pyplot as plt
import matplotlib.image as mpimg
import hashlib
import ipywidgets as widgets
from ipywidgets import FileUpload
from IPython.display import display
from PIL import Image
import os
import resources
from resources.utils import *

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

## Question 1: Getting started

In the cell below define the variable `x = "Hello, World!"`. 

```python
# ANSWER CELL
x = ...
print(x)
```

```python
grader.check("q1.0")
```

### You're done with Lab 00!

**Important submission information:** After completing the assignment, click on the Save icon from the Tool Bar &nbsp; <i class="fa fa-save" style="font-size:16px;"></i> &nbsp;. After saving your notebook, **run the cell with** `grader.check_all()` and confirm that you pass the same tests as in the notebook. Then, **run the final cell** `grader.export()` and click the link to download the zip file. Then, go to Gradescope and submit the zip file to the corresponding assignment. 

**Once you have submitted, stay on the Gradescope page to confirm that you pass the same tests as in the notebook.**

```python
%matplotlib inline
img = mpimg.imread('resources/animal.jpg')
imgplot = plt.imshow(img)
imgplot.axes.get_xaxis().set_visible(False)
imgplot.axes.get_yaxis().set_visible(False)
print("Congratulations on finishing lab 00!")
plt.show()
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Make sure you submit the zip file to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1.0 (points: 5)
```
>>> assert get_hash(x.upper()) == '1f9bd06f0b9c061cb1079a2bc30ded89'
```
